# Aufgabe 1 

Im Rahmen der operativen Planung in einem Handelsunternehmen soll auf der Basis unterschiedlicher
Faktoren wie z.B. dem Absatz und dem Lagerbestand berechnet werden, ob eine Handelsware vom
Lieferanten beschafft werden muss. 

Im ersten Schritt soll die Beschaffung lediglich vom Absatz des Vortages abhängen. Hierfür sollen folgende 
Bedingungen gelten: 

- Eine Handelsware wird beschafft, wenn am Vortag der Absatz mindestens 100 Stück war. 
- Die Bestellmenge soll doppelt so groß wie der Absatz sein. 

Implementieren Sie nun die Funktion `bestellmenge_berechnen`. Die Funktion erhält den Parameter `absatz_vortag`. 
Der Parameter ist vom Typ Integer und enthält den Absatz des Vortags in Stück. Als Ergebnis soll 
die Variabe `bestellmenge` berechnet werden. Die Variable `bestellmenge` gibt an wie viele Stücke einer 
Handelsware beschafft werden soll. Soll die Handelsware aufgrund der oben genannten Bedingungen aktuell nicht
beschafft werden soll die `bestellmenge` auf den Wert `0` gesetzt werden. 

## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandeln Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

## Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

```python
def bestellmenge_berechnen(absatz_vortag):
    
    bestellmenge = 0
    ### START

    # An dieser Stelle müssen Sie die pasende Bedingung formulieren
    if # Bedingung:
        bestellmenge = absatz_vortag * 2
    ### ENDE 
    return bestellmenge

print("Wenn der Absatz am Vortag 500 Stück ist berechent Ihre Lösung eine Beschaffungsmenge von: ", bestellmenge_berechnen(500))
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


## Test 1 - Korrekte Bedingung
Die Bestellmenge (`bestellmenge`) soll größer als 0 sein wenn von dem Produkt am Vortag (`absatz_vortag`) mindestens 100 Stück verkauft wurde. Ansonsten soll sie gleich 0 sein.

```python
import sys
sys.path.append("../..")
from nose.tools import assert_equal, assert_not_equal
from utils import * 

try:
    absatz_vortag = 50
    bestellmenge = bestellmenge_berechnen(absatz_vortag)

    assert_equal(bestellmenge, 0)
    print("[OK] Test Bestellmenge = 0 bei Absatz < 100 Stück erfolgreich. ")
except AssertionError as e:
    print("[FEHLER] Die Bestellmenge soll 0 sein wenn der Absatz des Vortags", absatz_vortag, "Stück ist.")
    print("       ", e)
except Exception as e:
    print("[FEHLER]", friendly_exceptions.msg(e)) 

    
try:
    absatz_vortag = 500
    bestellmenge = bestellmenge_berechnen(absatz_vortag)
    assert_not_equal(bestellmenge, 0)
    print("[OK] Test Bestellmenge > 0 bei Absatz > 100 Stück erfolgreich. ")
except AssertionError as e:
    print("[FEHLER] Die Bestellmenge soll größer als 0 sein wenn der Absatz des Vortags", absatz_vortag, "ist.")
    print("       ", e)
except Exception as e:
    print("[FEHLER]", friendly_exceptions.msg(e)) 
```

## Test 2 - Beachtung der Randfälle
Die Bestellmenge (`bestellmenge`) soll größer als 0 sein wenn von dem Produkt am Vortag (absatz_vortag) **mindestens** 100 Stück verkauft wurde.
Das heißt, dass bei eimm Absatz von **genau** 100 Stück am Vortag die Bestellmenge 200 Stück sein soll.

```python
import sys
sys.path.append("../..")
from nose.tools import assert_equal, assert_not_equal
from utils import * 

absatz_vortag = 100
bestellmenge = bestellmenge_berechnen(absatz_vortag)
try:
    assert_equal(bestellmenge, 200)
    print("[OK] Test Bestellmenge = 200 bei Absatz = 100 Stück erfolgreich. ")
except AssertionError as e:
    print("[FEHLER] Die Bestellmenge soll 200 sein wenn der Absatz des Vortags", absatz_vortag, "ist.")
    print("       ", e)
except Exception as e:
    print("[FEHLER]", friendly_exceptions.msg(e)) 
```

## Test 3 - Korrekte Berechnung der Bestellmenge
Die Bestellmenge (`bestellmenge`) soll korrekt berechnet werden. Nämlich als Bestellmenge = 2 * Absatz des Vortags.

<!-- grade_id: cell-b7ba8a17f5efa24a, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

absatz_vortag = 345
bestellmenge = bestellmenge_berechnen(absatz_vortag)
try:
    assert_equal(bestellmenge, 690)
    print("Test Bestellmenge = 690 Stück bei Absatz 100 Stück erfolgreich.")
except AssertionError as e:
    print("Das Ergebnis soll 2 mal so groß sein wie \"absatz_vortag\".")
    print(e)
```