# Aufgabe 4

Es soll ein Programm zur automatischen Bewertung mehrerer Konten implementiert werden. So sollen aus einer Liste mehrerer Konten alle herausgefiltert werden, deren Kontostand negativ ist.  

Implementieren Sie nun die Funktion `analyse_accounts`. Diese Funktion erhält als Parameter eine Liste `accounts`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `accounts` eine Liste bestehend aus dem Kontonamen als erstem Wert, gefolgt von dem Kontostand als zweitem Wert. Die Liste sieht also in etwa so aus: `[["Kasse", +330], ["Immobilien", +25000], ["Fuhrpark", -16500]]`.  

Als Ergebnis soll eine Liste mit nur den negativen Einträgen zurückgegeben werden. Ergebnis für obiges Beispiel wäre also: `[["Fuhrpark", -16500]]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-302c3da36eb63d24, points: 0, grade: false, solution: true -->
```python
def analyse_accounts(accounts):
    
    negative_accounts = []
    for account in accounts:
        if #Bedingung wann ein Konto zu der Liste, die später zurückgegeben wird, hinzugefügt werden soll:
            negative_accounts.append(account)
    
    return negative_accounts

print("Bei \"accounts = [[\"Kasse\", +330], [\"Immobilien\", +25000]]\" liefert Ihre Funktion:", analyse_accounts([["Kasse", +330], ["Immobilien", -25000]]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-888b94242172b887, points: 0, grade: true, solution: false -->
```python
try:
    analyse_accounts([["Kasse", +330], ["Immobilien", +25000]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Richtige Filterung
Es sollen die Konten herausgefiltert werden, bei denen der Kontostand negativ ist.

<!-- grade_id: cell-cdd05810d4fbdcc6, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

accounts = [["Kasse", +330], ["Immobilien", +25000], ["Fuhrpark", -16500]]
negative_accounts = analyse_accounts(accounts)
try:
    assert_equal(negative_accounts, [["Fuhrpark", -16500]])
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Konto \"Fuhrpark\" sollte als einziges Konto zurückgegeben werden")
    print(e)
    

accounts = [["Kasse", -330], ["Immobilien", -25000], ["Fuhrpark", -16500]]
negative_accounts = analyse_accounts(accounts)
try:
    assert_equal(negative_accounts, [["Kasse", -330], ["Immobilien", -25000], ["Fuhrpark", -16500]])
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Die Konten \"Kasse, Immobilien und Fuhrpark\" sollten alle zurückgegeben werden")
    print(e)
    

accounts = [["Kasse", -330], ["Immobilien", +25000], ["Fuhrpark", +16500], ["Lager", -200]]
negative_accounts = analyse_accounts(accounts)
try:
    assert_equal(negative_accounts, [["Kasse", -330], ["Lager", -200]])
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Die Konten \"Kasse und Lager\" sollten zurückgegeben werden")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```