Aufgabe 4 - Aktien filtern
------------
Gegeben ist eine Liste mit Aktienkursen.  
Jedes Element dieser Liste ist jedoch selbst eine Liste. Diese enthält als ersten Wert den Namen der Aktie, als zweiten die Veränderung zum Vortag und als dritten den Wert einer Aktie.  
Nun wollen Sie alle gut im Wert gestiegenen Aktien herausfiltern um diese zu verkaufen.  
Erstellen Sie dazu eine neue Liste, wieder bestehend aus Listen, in welcher Sie alle Aktienkurse speichern mit einer Veränderung zum Vortag größer/gleich 1.  

<!-- grade_id: cell-ae387673cf8e32fa, points: 0, grade: false, solution: true -->
```python
def filter_aktien(aktien):
    
    aktien_zu_verkaufen = []
    ###START
    for #Schleife:
        if # Bedingung um aktie in die Liste mit aufzunehmen:
            #Aktie in die neue Liste aufnehmen
    ###ENDE
     
    return aktien_zu_verkaufen

print("Bei \"aktien = ['SAP', -0.3, 100.0], ['MSFT', +1.5, 180.4]\" liefert Ihre Funktion:", filter_aktien([['SAP', -0.3, 100.0], ['MSFT', +1.5, 180.4]]))
```

   ***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-5ae06540cdad36c5, points: 0, grade: true, solution: false -->
```python
try:
    filter_aktien([['SAP', -0.3, 100.0], ['MSFT', +1.5, 180.4]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Filterung der meisten Aktien
Die von der Funktion `filter_aktien` zurückgegebene Liste `aktien_zu_verkaufen` muss die korrekten Aktienkurse enthalten.

<!-- grade_id: cell-5930e3ee07100aa3, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

shares = [['SAP', -0.3, 100.0], ['MSFT', +1.5, 180.4], ['APPL', +0.5, 480.4], ['AMZN', +1.23, 450.8], ['MCD', -1.2, 350]]
try:
    expected_output = [['MSFT', 1.5, 180.4], ['AMZN', 1.23, 450.8]]
    assert_equal(filter_aktien(shares), expected_output)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Die zurückgegebenen Aktienkurse sind nicht korrekt.")
    print(e)
```

## Test 2 - Einbeziehung der Aktien mit Veränderungen genau gleich 1
Die von der Funktion `filter_aktien` zurückgegebene Liste `aktien_zu_verkaufen` soll nicht nur alle Aktienkurse mit einer Veränderung zum Vortag größer als 1 speichern sondern auch Aktienkurse bei denen die Veränderung genau gleich 1 ist.

<!-- grade_id: cell-e761db6779b21c92, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

shares = [['MSFT', +1, 180.4]]
try:
    expected_output = [['MSFT', +1, 180.4]]
    assert_equal(filter_aktien(shares), expected_output)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Die Liste bezieht keine Aktien ein, bei denen die Veränderung genau gleich 1 ist.")
    print(e)
```