# Aufgabe 3

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Analyse von Materialeingängen und Materialabgängen implementiert werden. So soll aus einer großen Liste von Eingängen und Abgängen erkannt werden wie der neue Materialbestand ist. Außerdem soll die durchschnittliche Höhe der Materialbewegungen berechnet werden, sowie der größte Ein- oder Abgang bestimmt werden.

Implementieren Sie nun die Funktion `analyse_bestand`. Diese Funktion erhält als Parameter eine Liste `eingehend_ausgehend`, welche in Form von Integer-Werten die Eingänge und Abgänge auf das Konto angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis soll der Materialbestand(`bestand`) als Integer-Wert zurückgegeben werden.  
Die durchschnittliche Höhe(`durchschnitt`) der Eingänge und Abgänge eines Materials.  
Die höchste Transaktion (`hoechste_transaktion`). Hierbei sollen sowohl Abgänge als auch Eingänge berücksichtigt werden.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-903a8769bae67c57, points: 0, grade: false, solution: true -->
```python
def analyse_bestand(eingehend_ausgehend):
    
    bestand = 0
    hoechste_transaktion = 0
    for transaktion in eingehend_ausgehend:
        ###START
        # Berechne mit jedem Durchlauf einen neuen Materialbestand
        # Prüfe außerdem ob die Transaktion die höchste ist
        ###ENDE
    durchschnitt = bestand / len(eingehend_ausgehend)
    
    return bestand, durchschnitt, hoechste_transaktion

print("Bei \"eingehend_ausgehend = [20, 30, 40]\" liefert Ihre Funktion:", analyse_bestand([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-5bb0046f57db8819, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bestand([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung des Kontostands
Der Materialbestand (`bestand`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-decd091a2b317c36, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = analyse_bestand(incoming_outgoing)
try:
    assert_equal(stock[0], 472)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = analyse_bestand(incoming_outgoing)
try:
    assert_equal(stock[0], 88888889)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [-12, +20]
stock = analyse_bestand(incoming_outgoing)
try:
    assert_equal(stock[0], 8)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

## Test 2 - Korrekte Berechnung des Durchschnitts
Der Durchschnitt (`durchschnitt`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-945d1aa59d6ae4d8, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+300, -15, +230, -123, +987]
average = analyse_bestand(incoming_outgoing)
try:
    assert_equal(average[1], 275.8)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+200, -200, +200, -200]
average = analyse_bestand(incoming_outgoing)
try:
    assert_equal(average[1], 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+15]
average = analyse_bestand(incoming_outgoing)
try:
    assert_equal(average[1], 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

## Test 3 - Korrekte Bestimmung der höchsten Materialtransaktion
Die höchste Warenbewegung (`hoechste_transaktion`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch Materialabgänge zu berücksichtigen.

<!-- grade_id: cell-138b09255c4e3865, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+100, -26, -23, +210, +234]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction[2], 234)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+200, -200, +200, -200]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction[2], 200)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+15]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction[2], 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+300, -500, -23, +210, +234]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction[2], -500)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```