# Aufgabe 3

Es soll ein Programm zur automatischen Analyse von Kontobewegungen implementiert werden. So soll aus einer großen Liste von Zugängen und Abgängen erkannt werden wie der neue Kontostand ist. Außerdem soll die durchschnittliche Höhe der Kontobewegungen berechnet werden, sowie der größte Zu- oder Abgang bestimmt werden.

Implementieren Sie nun die Funktion `analyse_bewegungen`. Diese Funktion erhält als Parameter eine Liste `bewegungen`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).  
Als Ergebnis soll der `Kontostand` als Float-Wert zurückgegeben werden ("kontostand").  
Die `durchschnittliche Höhe` einer Kontobewegung ("durchschnitt").  
Der `höchste Transfer` ("hoechste_bewegung"). Hierbei sollen sowohl negative als auch positive Transfers berücksichtigt werden.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-9debd1ae65757e87, points: 0, grade: false, solution: true -->
```python
def analyse_bewegungen(bewegungen):
    
    kontostand = 0
    hoechste_bewegung = 0
    for bewegung in bewegungen:
    ###START
        # Berechne den Kontostand mit jeder neuen Buchung neu
        if bewegungen ** 2 > hoechste_bewegung ** 2:
            hoechste_bewegung = bewegung
    
    # Berechne den Durchschnitt
    ###ENDE
    
    return kontostand, durchschnitt, hoechste_bewegung

print("Bei \"bewegungen = [20, 30, 40]\" liefert Ihre Funktion:", analyse_bewegungen([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-891aa9f193fdb110, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bewegungen([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung des Kontostands
Der Kontostand (`kontostand`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-55e02cc4dbc5a4b1, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+519.55, -250.00, -19.37, -12.34, +234.56]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[0], 472.40)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+99999999.99, -1234567.89, -9876543.22]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[0], 88888888.88)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [-12.34, +20.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[0], 7.66)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

## Test 2 - Korrekte Berechnung des Durchschnitts
Der Durchschnitt (`durchschnitt`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-0565a8614ee86169, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+300.00, -250.00, -23.00, +210.00, +234.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[1], 94.2)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+200.00, -200.00, +200.00, -200.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[1], 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+15.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[1], 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

## Test 3 - Korrekte Bestimmung der höchsten Bankbewegung
Die höchste Bankbewegung (`hoechste_bewegung`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-34c24ceab6d073ba, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+300.00, -250.00, -23.00, +210.00, +234.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[2], 300)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+200.00, -200.00, +200.00, -200.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[2], 200)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+15.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[2], 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+300.00, -500.00, -23.00, +210.00, +234.00]
balance = analyse_bewegungen(transfers)
try:
    assert_equal(balance[2], -500)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
    print(e)
```