# Aufgabe 5

Im Rahmen der Bestellungsabwicklung soll ein Programm zur automatischen Bewertung mehrerer Bestellungen implementiert werden. So soll anhand einer Liste von Bestellungen und einer (separaten) Liste von Zahlungen erkannt werden, welche Bestellungen noch nicht bezahlt wurden.  

Implementieren Sie nun die Funktion `analyse_bestellungen`. Diese Funktion erhält als Parameter eine Liste `bestellungen`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `bestellungen` eine Liste bestehend aus der Bestellnummer als erstem Wert, gefolgt von einer Artikelbezeichnung. Die Liste sieht also in etwa so aus: `[[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], ["10003", "Früchtemüsli"]]`.   
Zusätzlich gibt es als zweiten Parameter eine weitere Liste `bezahlte`. Diese Liste beinhält alle Bestellnummern, welche bezahlt wurden. Also in etwa so: `[10001, 10003]`.  
Als Ergebnis sollen alle Bestellungen zurückgegeben werden, welche noch nicht bezahlt wurden. Für die obigen Beispiele wäre die Ergebnis-Liste also: `[[10002, "Blaubeermüsli"]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-65b62e49dfcb7d3c, points: 0, grade: false, solution: true -->
```python
def analyse_bestellungen(bestellungen, bezahlte):
    
    unbezahlte_bestellungen = bestellungen

    for bestellung in unbezahlte_bestellungen: 
        for bezahlt in bezahlte:
            ###START
            if #Bedingung wann eine Bestellung aus der Liste "unbezahlte_bestellungen" entfernt werden soll:
            ###ENDE
                unbezahlte_bestellungen.remove(bestellung)                

    return unbezahlte_bestellungen

print("Bei \"bestellungen = [[10001, \"Erdbeermüsli\"], [10002, \"Blaubeermüsli\"]]\" und \"payments = [10001, 10003]\" liefert Ihre Funktion:", analyse_bestellungen([[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"]], [10001, 10003]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-8cccd9a10c5cdf35, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bestellungen([[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"]], [10001, 10003])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Filterung der Bestellungen
Die noch *nicht* bezahlten Bestellungen sollen zurückgegeben werden.

<!-- grade_id: cell-237bf4f739287ceb, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"]]
payments = [10001, 10003]
paid_orders = analyse_bestellungen(orders, payments)
try:
    assert_equal(paid_orders, [[10002, "Blaubeermüsli"]])
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Die Bestellung \"10002 Blaubeermüsli\" soll zurückgegeben werden.")
    print(e)
    

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"], [10004, "Früchtemüsli"]]
payments = [10001, 10003]
paid_orders = analyse_bestellungen(orders, payments)
try:
    assert_equal(paid_orders, [[10002, "Blaubeermüsli"], [10004, "Früchtemüsli"]])
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Die Bestellungen \"10002 Blaubeermüsli\" und \"10004 Früchtemüsli\" sollen zurückgegeben werden.")
    print(e)
    

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"], [10004, "Früchtemüsli"]]
payments = [10004]
paid_orders = analyse_bestellungen(orders, payments)
try:
    assert_equal(paid_orders, [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"]])
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Die Bestellungen \"10001 Erdbeermüsli\", \"10002 Blaubeermüsli\" und \"10003 Früchtemüsli\" sollen zurückgegeben werden.")
    print(e)
```