# Aufgabe 1

Für Ihren Online-Shop möchten Sie ein Programm implementieren, welches einen Gutscheincode überprüft und gegebenenfalls anwendet wenn dieser korrekt ist.  

Implementieren Sie nun die Funktion `validiere_coupon`. Diese Funktion soll nach einer Eingabe verlangen. Stimmt diese Eingabe *nicht* mit dem aktuellen Couponcode "Winter19" überein, so soll die Eingabe immer weiter wiederholt werden, bis die richtige Eingabe erfolgt. Das soll folgendermaßen aussehen:  

Coupon Code: __12345678__  
Coupon Code: __passwort__  
Coupon Code: __winter19__  
Coupon Code: __Winter19__  
Coupon angewendet.




### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `While-Schleifen` in Python finden Sie [hier](../../50_loops/while_loop.ipynb).

<!-- grade_id: cell-1e0d00fceb6de555, points: 0, grade: false, solution: true -->
```python
def validiere_coupon():
    
    ###START
    
    ###ENDE
    
    return

validiere_coupon()
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-fe796c6b0d257efa, points: 0, grade: true, solution: false -->
```python
try:
    validiere_coupon()()
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Abfrage
Das Programm soll wie oben beschrieben funktionieren.

__Hinweis__: Diese Tests sind sehr anfällig für kleinste Abweichungen. Achten Sie bitte darauf, dass Ihr Programm *exakt* so aussieht wie in der Aufgabenstellung als Beispiel dargestellt.  
Achten Sie bitte auch darauf den Text "Coupon Code: " mittels input auszugeben (also `input("Coupon Code: ")`). Bitte verwenden Sie hier kein print().

<!-- grade_id: cell-b1f69e1f002fb70f, points: 0, grade: true, solution: false -->
```python
from unittest import mock
from nose.tools import assert_equal
import unittest
import io

with mock.patch('builtins.input', side_effect=["Winter19"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()()

    print("Teste Winter19 als erste Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Coupon angewendet.\n', "Auch wenn Winter19 die erste Eingabe ist, sollte das Programm schon beendet werden")
    print("Test 1 erfolgreich.")
    
with mock.patch('builtins.input', side_effect=["a", "199", "Winter19"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()()

    print("Teste Winter19 als dritte Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Coupon angewendet.\n', "Das Programm soll so lange eine Eingabe verlangen bis diese Winter19 ist.")
    print("Test 2 erfolgreich.")

```