# Aufgabe 1

Es soll ein Programm zur automatischen Berechnung eines Kontostandes anhand von Zugängen und Abgängen auf ein Konto implementiert werden.

Implementieren Sie nun die Funktion `berechne_kontostand`. Diese Funktion erhält als Parameter eine Liste `ueberweisungen`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).  
Als Ergebnis soll der Kontostand in Form eines Float-Wertes zurückgegeben werden ("kontostand").


### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-e9702910f65c178d, points: 0, grade: false, solution: true -->
```python
def berechne_kontostand(ueberweisungen):
    ###START
    #instanziiere die Variable kontostand
    for ueberweisung in ueberweisungen:
        # ergänze die Variable kontostand um eine zubuchung oder abbuchung
    ###ENDE
    return kontostand

print("Bei \"transfers = [20, 30, 40]\" liefert Ihre Funktion:", berechne_kontostands([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-a8425fbea7f36c46, points: 0, grade: true, solution: false -->
```python
try:
    berechne_kontostand([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Kontostand (`balance`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-004216048770b354, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+519.55, -250.00, -19.37, -12.34, +234.56]
balance = berechne_kontostand(transfers)
try:
    assert_equal(balance, 472.40)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    
    
transfers = [+99999999.99, -1234567.89, -9876543.22]
balance = berechne_kontostand(transfers)
try:
    assert_equal(balance, 88888888.88)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    
    
transfers = [-12.34, +20.00]
balance = berechne_kontostand(transfers)
try:
    assert_equal(balance, 7.66)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```