# Aufgabe 1

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Berechnung der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden.   

Implementieren Sie nun die Funktion `calculate_taxes`. Diese Funktion erhält als Parameter eine Liste `sales`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).  
Als Ergebnis soll der gesamte Betrag, welcher durch die Umsatzsteuer eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden ("taxes").   

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-baa58e3160f441e1, points: 0, grade: false, solution: true -->
```python
def calculate_taxes(sales):
    
    ###BEGIN SOLUTION
    taxes = 0
    for sale in sales:
        taxes += (sale * 0.19)
    
    ###END SOLUTION
    return taxes

print("Bei \"sales = [20, 30, 40]\" liefert Ihre Funktion:", calculate_taxes([20, 30, 40]))
```

Output:
```
Bei "sales = [20, 30, 40]" liefert Ihre Funktion: 17.1
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-203187fb7d4cf33c, points: 0, grade: true, solution: false -->
```python
try:
    calculate_taxes([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Betrag (`taxes`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-9e876bfe5063e32e, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300, 250.45, 19.78, 123.56, 50.23]
taxes = calculate_taxes(sales)
try:
    assert_equal(taxes, 141.3638)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print()
    print(e)
    

sales = [99999999, 1234567, 9876543]
taxes = calculate_taxes(sales)
try:
    assert_equal(taxes, 21111110.71)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [12, 20]
taxes = calculate_taxes(sales)
try:
    assert_equal(taxes, 6.08)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```