# Aufgabe 2

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Analyse der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden. So soll aus einer Liste mit vielen Einnahmen herausgefiltert werden, welches die höchste eingenommene Umsatzsteuer ist.  

Implementieren Sie nun die Funktion `analyse_steuern`. Diese Funktion erhält als Parameter eine Liste `verkaeufe`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).  
Als Ergebnis soll der höchste Betrag, welcher durch die Umsatzsteuer an einem Artikel eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden ("hoechste_steuern").   

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-a5f281ffaac0dc9e, points: 0, grade: false, solution: true -->
```python
def analyse_steuern(verkaeufe):
    
    hoechste_steuern = 0
    for verkauf in verkaeufe:
        if (verkauf * 0.19) > hoechste_steuern:
            ###START
            # ersetze den Wert der hoechste_steuern durch die neue Umsatzsteuer des Verkaufs
            ###ENDE
    return hoechste_steuern

print("Bei \"verkaeufe = [20, 30, 40]\" liefert Ihre Funktion:", analyse_steuern([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-6a6176b58b497b62, points: 0, grade: true, solution: false -->
```python
try:
    analyse_steuern([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bestimmung der höchsten EIngenommenen Umsatzsteuer
Der Betrag (`hoechste_steuern`) soll für jede Liste korrekt herausgefiltert werden.

<!-- grade_id: cell-e87fcf79985d5938, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300, 250.45, 19.78, 123.56, 50.23]
highest_tax = analyse_steuern(sales)
try:
    assert_equal(highest_tax, 57)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [99999999, 1234567, 9876543]
highest_tax = analyse_steuern(sales)
try:
    assert_equal(highest_tax, 18999999.81)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [12, 20]
highest_tax = analyse_steuern(sales)
try:
    assert_equal(highest_tax, 3.8)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [20]
highest_tax = analyse_steuern(sales)
try:
    assert_equal(highest_tax, 3.8)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```