# Aufgabe 4

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Bewertung mehrerer Materialien implementiert werden. So sollen aus einer Liste mehrerer Materialien alle herausgefiltert werden, deren Lagerbestand unter dem Meldebestand von 200 liegt, damit diese neu bestellt werden können.  

Implementieren Sie nun die Funktion `analyse_bestand`. Diese Funktion erhält als Parameter eine Liste `materialien`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `materialien` eine Liste bestehend aus dem Materialnamen als erstem Wert, gefolgt von dem Materialbestand als zweitem Wert. Die Liste sieht also in etwa so aus: `[["Schrauben", 330], ["Nägel", 2500], ["Muttern", 150]]`.  

Als Ergebnis soll eine Liste mit nur den Materialien unter dem Meldebestand von 200Stück zurückgegeben werden (`neuzubestellen`). Ergebnis für obiges Beispiel wäre also: `[["Muttern", 150]]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-15704bf10445fdef, points: 0, grade: false, solution: true -->
```python
def analyse_bestand(materialien):
    
    neuzubestellen = []
    for material in materialien:
        ###START
        # Prüfe ob das Material einen Bestand von unter 200 hat und füge es gegebenenfalls zur Liste "neuzubestellen" hinzu
        ###ENDE
    return neuzubestellen

print("Bei \"materialien = [[\"Schrauben\", 330], [\"Nägel\", 2500]]\" liefert Ihre Funktion:", analyse_bestand([["Schrauben", 330], ["Nägel", 150]]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-a74fc5bb73d9004c, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bestand([["Schrauben", 330], ["Nägel", 2500]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Richtige Filterung
Es sollen die Materialien herausgefiltert werden, bei denen der Bestand unter 200 Stück liegt.

<!-- grade_id: cell-724ad39562732b91, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

materials = [["Schrauben", 330], ["Nägel", 2500], ["Muttern", 150]]
to_reorder = analyse_bestand(materials)
try:
    assert_equal(to_reorder, [["Muttern", 150]])
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Material \"Muttern\" sollte als einziges Material zurückgegeben werden")
    print(e)
    

materials = [["Gewinde", 230], ["Knopfzellen", 120]]
to_reorder = analyse_bestand(materials)
try:
    assert_equal(to_reorder, [["Knopfzellen", 120]])
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Material \"Knopfzellen\" sollte als einziges Material zurückgegeben werden")
    print(e)
    

materials = [["Mandeln", 150], ["Haselnüsse", 2500], ["Erdnüsse", 300], ["Haselnüsse", 180]]
to_reorder = analyse_bestand(materials)
try:
    assert_equal(to_reorder, [["Mandeln", 150], ["Haselnüsse", 180]])
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Die Materialien \"Mandeln\" und \"Haselnüsse\" sollten als zurückgegeben werden")
    print(e)
```