# Aufgabe 4

Im Rahmen der Bestellungsabwicklung soll ein Programm zur automatischen Bewertung mehrerer Bestellungen implementiert werden. So sollen aus einer Liste mehrerer Bestellungen alle herausgefiltert werden, die schon bezahlt wurden.

Implementieren Sie nun die Funktion `analyse_bestellungen`. Diese Funktion erhält als Parameter eine Liste `bestellungen`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `bestellungen` eine Liste bestehend aus der Bestellungsnummer als erstem Wert, gefolgt von der Artikelbezeichnung als zweitem Wert, sowie als drittes einem Boolean-Wert der speichert ob die Bestellung bezahlt wurde (in diesem Fall ist der Wert True). Die Liste sieht also in etwa so aus: `[[76576, "Regal", True], [76577, "Tisch", False], [76578, "Bett", True]]`.  

Als Ergebnis soll eine Liste mit nur den Bestellungen, die schon bezahlt wurden, zurückgegeben werden (`bezahlte_bestellungen`). Ergebnis für obiges Beispiel wäre also: `[[76576, "Regal", True], [76578, "Bett", True]]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-e99d1d9aa40f8383, points: 0, grade: false, solution: true -->
```python
def analyse_bestellungen(bestellungen):
    
    bezahlte_bestellungen = []
    for bestellung in bestellungen:
        ###START
        # prüfe ob die Bestellung bezahlt ist und füge sie gegebenenfalls zu "bezahlte_bestellungen" hinzu
        ###ENDE   
    return bezahlte_bestellungen

print("Bei \"bestellungen = [[10001, \"Esstisch\", False], [10002, \"Sofa\", True]]\" liefert Ihre Funktion:", analyse_bestellungen([[10001, "Esstisch", False], [10002, "Sofa", True]]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-c4b4edfe13551af1, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bestellungen([[10001, "Esstisch", False], [10002, "Sofa", False]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Richtige Filterung
Es sollen die Bestellungen herausgefiltert werden, die bereits bezahlt wurden. Das heißt der dritte Wert ist "True".

<!-- grade_id: cell-ffc6c3a680fd5227, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

orders = [[10001, "Esstisch", False], [10002, "Sofa", False], [10003, "Schreibtisch", True]]
paid_orders = analyse_bestellungen(orders)
try:
    assert_equal(paid_orders, [[10003, "Schreibtisch", True]])
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Material \"Muttern\" sollte als einziges Material zurückgegeben werden")
    print(e)
    

orders = [[10001, "iPhone X", True], [10002, "iPad 2018", False], [10003, "Apple Watch", False]]
paid_orders = analyse_bestellungen(orders)
try:
    assert_equal(paid_orders, [[10001, "iPhone X", True]])
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Material \"Muttern\" sollte als einziges Material zurückgegeben werden")
    print(e)
    

orders = [[10001, "Esstisch", True], [10002, "Sofa", False], [10003, "Schublade", False], [10004, "Schreibtisch", True]]
paid_orders = analyse_bestellungen(orders)
try:
    assert_equal(paid_orders, [[10001, "Esstisch", True], [10004, "Schreibtisch", True]])
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Material \"Muttern\" sollte als einziges Material zurückgegeben werden")
    print(e)
```