Aufgabe 3 - Produktionsmengenplanung
------------
Gegeben sind die Absatzprognosen sowie Ziellagerbestände eines Produktes für die nächsten 12 Monate in gleichnamigen Listen.
Erstellen Sie eine neue Liste, in der Sie die Produktionsmengen für die jeweiligen Monate speichern.  
Beachten Sie die Lagerbestände aus dem jeweils vorhergehenden Monat in den nächsten Monat zu übernehmen.

<!-- grade_id: cell-073db0e144d454aa, points: 0, grade: false, solution: true -->
```python
def erstelle_produktions_plan(verkaufs_plan, ziel_bestand):
    
    produktions_plan = []
    vorheriger_bestand = 0
    for monat in range(len(verkaufs_plan)):
        ###START
        # Berechnung der Produktionsmengen für den aktuellen Monat
        ###ENDE
        vorheriger_bestand = ziel_bestand[monat]
        
    return produktions_plan

print("Bei \"verkaufs_plan = [20, 30, 40]\" und \"ziel_bestand = [50, 60, 65]\" liefert Ihre Funktion:", erstelle_produktions_plan([20, 30, 40], [50, 60, 65]))
```

   ***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-42ba2270619bd540, points: 0, grade: true, solution: false -->
```python
try:
    erstelle_produktions_plan([0, 0, 0, 0], [0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Länge der Ergebnisliste
Das von der Funktion `erstelle_produktions_plan` zurückgegebene Ergebnis  `produktions_plan` ist eine Liste.
Diese Liste muss genau so viele Einträge enthalten wie
die übergebenen Listen `verkaufs_plan` und `ziel_bestand`. In diesem Test hat die übergebene Liste die Länge 12.

<!-- grade_id: cell-dd7d08596810b7d2, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

sales_plan = [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
target_stock = [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
production_plan = erstelle_produktions_plan(sales_plan, target_stock)
try:
    assert_is_instance(production_plan, list)
    assert_equal(len(production_plan), 12)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Die Produktionsmengen sollen als Liste ausgegeben werden.\nDie zurückgegebene Liste muss 12 Einträge enthalten.")
    print(e)
```

## Test 2 - Korrekte Berechnung der Produktionsmengen
Die von der Funktion `erstelle_produktions_plan` zurückgegebene Liste `produktions_plan` muss die korrekten Produktionsmengen in
Abhängigkeit von den übergebenen Absatzmengen und Ziellagerbeständen enthalten.

<!-- grade_id: cell-c25447d271049c6a, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

sales_plan = [570, 500, 520, 500, 600, 550, 530, 520, 530, 590, 600, 590]
target_stock = [100, 100, 150, 100, 200, 100, 150, 100, 150, 100, 100, 50]
try:
    expected_production_plan = [670, 500, 570, 450, 700, 450, 580, 470, 580, 540, 600, 540]
    assert_equal(erstelle_produktions_plan(sales_plan, target_stock), expected_production_plan)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Die berechneten Produktionsmengen sind nicht korrekt.")
    print(e)
```