# Aufgabe 1

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Berechnung eines Materialbestandes, anhand von Eingängen und Abgängen eines Materials in einer Lagerhalle, implementiert werden.  

Implementieren Sie nun die Funktion `calculate_stock`. Diese Funktion erhält als Parameter eine Liste `incoming_outgoing`, welche in Form von Integer-Werten die Materialeingänge und Abgänge in die Lagerhalle angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis soll der Materialbestand in Form eines Integer-Wertes zurückgegeben werden ("stock").


### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-080ae4f875168748, points: 0, grade: false, solution: true -->
```python
def calculate_stock(incoming_outgoing):
    
    stock = 0
    for entry in incoming_outgoing:
        # addiere oder subtrahiere den Betrag des einzelnen Wareneingangs oder Abgangs zur Variable "stock"
    
    return stock

print("Bei \"incoming_outgoing = [20, 30, 40]\" liefert Ihre Funktion:", calculate_stock([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-9563226530e72506, points: 0, grade: true, solution: false -->
```python
try:
    calculate_stock([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Materialbestand (`stock`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-e412b7ff90a78f87, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 472)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 88888889)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [-12, +20]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 8)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```