Aufgabe 5 - Aktienkurs darstellen
------------
Gegeben ist eine Liste mit Werten einer Aktie über den Zeitraum von einem Jahr. Die Liste beinhält 12 Werte, also für die Monate Januar bis Dezember. Diese Werte sind vereinfacht auf 50 gerundet wobei der kleinste Wert 50 ist und der höchste Wert 300.  
Ihre Aufgabe ist es diese Werte in einem Diagramm grafisch darzustellen.  
Erstellen Sie dazu eine y-Achse mit den Werten von 50 bis 300 und eine x-Achse mit den Monaten Januar bis Dezember. 
Tragen Sie für jeden Monat den Aktienwert mit einem "x" ein.  

Für eine Liste `100, 150, 200, 250, 200, 250, 200, 100, 150, 250, 200, 200`,  
sähe das Diagramm dann so aus:  
`
300                                     
250           x     x           x       
200        x     x     x           x  x 
150     x                    x          
100  x                    x             
50                                      
     J  F  M  A  M  J  J  A  S  O  N  D` 


<!-- grade_id: cell-ae387673cf8e32fa, points: 0, grade: false, solution: true -->
```python
def aktien_visualisieren(aktien):
    
    for wert in range(300, 49, -50):
        print(wert, ende="")
        ###START
        # zweite For-Schleife zum durchlaufen der Monate
        ###ENDE
        print()
    print("     J  F  M  A  M  J  J  A  S  O  N  D")
            
    return

aktien_visualisieren([100, 150, 200, 250, 200, 250, 200, 100, 150, 250, 200, 200])
```

   ***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.

Die für diese Aufgabe verwendeten Tests sind schon für kleinste Abweichungen von der Musterlösung anfällig. Achten Sie bitte darauf dass ihre Ausgabe exakt so aussieht wie hier dargestellt.  
- Tipps:
  - Der Abstand vor bzw. zwischen zwei "x" ist immer 2 Leerzeichen groß.  
  - Der Zeilenumbruch erfolgt immer an der selben Stelle, unabhängig davon wo das "x" in dieser Zeile steht
  - Vor dem "J" für Januar stehen 5 Leerzeichen
  - zwischen den Buchstaben für die Monate stehen jeweils 2 Leerzeichen
  - Nach dem "D" für Dezember erfolgt ein Zeilenumbruch

<!-- grade_id: cell-3c1b47bde9f320b1, points: 0, grade: true, solution: false -->
```python
try:
    aktien_visualisieren([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Darstellung einer Liste
Für die Liste `100, 150, 200, 250, 200, 250, 200, 100, 150, 250, 200, 200`,  
soll das Diagramm so aussehen:  
`
300                                     
250           x     x           x       
200        x     x     x           x  x 
150     x                    x          
100  x                    x             
50                                      
     J  F  M  A  M  J  J  A  S  O  N  D` 


<!-- grade_id: cell-5930e3ee07100aa3, points: 1, grade: true, solution: false -->
```python
from unittest import mock
import unittest
import io
from nose.tools import assert_equal, assert_not_equal


with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    shares = [100, 150, 200, 250, 200, 250, 200, 100, 150, 250, 200, 200]
    aktien_visualisieren(shares)

assert_equal(fake_stdout.getvalue(), '300                                    \n250           x     x           x      \n200        x     x     x           x  x\n150     x                    x         \n100  x                    x            \n50                                    \n     J  F  M  A  M  J  J  A  S  O  N  D\n', "Die Ausgabe ist nicht korrekt")
print("Test 1 erfolgreich")
```

## Test 2 - Korrekte Darstellung einer Liste
Für die Liste `100, 300, 150, 250, 200, 250, 300, 100, 150, 200, 250, 300`,  
soll das Diagramm so aussehen:  
`
300     x              x              x 
250           x     x              x    
200              x              x       
150        x                 x          
100  x                    x             
50                                      
     J  F  M  A  M  J  J  A  S  O  N  D`


<!-- grade_id: cell-a569658825146cde, points: 1, grade: true, solution: false -->
```python
from unittest import mock
import unittest
import io
from nose.tools import assert_equal, assert_not_equal


with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    shares = [100, 300, 150, 250, 200, 250, 300, 100, 150, 200, 250, 300]
    aktien_visualisieren(shares)

assert_equal(fake_stdout.getvalue(), '300     x              x              x\n250           x     x              x   \n200              x              x      \n150        x                 x         \n100  x                    x            \n50                                    \n     J  F  M  A  M  J  J  A  S  O  N  D\n', "Die Ausgabe ist nicht korrekt")
print("Test 2 erfolgreich")
```