# Aufgabe 3

Um Ihren Online-Shop vor Hackerangriffen zu schützen wollen Sie ein Programm implementieren, welches beim betreten der Webseite überprüft ob der Besucher ein Mensch ist. Bekannt sind diese Tests als [Captcha](https://en.wikipedia.org/wiki/CAPTCHA).  
Versuchen Sie in diesem Schritt ein vereinfachtes Captcha zu programmieren, welches vom Benutzer eine Primzahl als Eingabe erwartet. 

Implementieren Sie nun die Funktion `captcha`. Diese Funktion soll nach einer Eingabe verlangen. Prüfen Sie zunächst ob diese Eingabe eine Zahl ist. Ist dies nicht der Fall so soll das Programm erst gar nicht weiter ausgeführt werden.   
Prüfen Sie danach ob die Zahl eine Primzahl ist. Ist diese Eingabe keine Primzahl, so soll die Eingabe immer weiter wiederholt werden, bis die richtige Eingabe erfolgt. Das soll folgendermaßen aussehen:  

Fall 1:  
Geben Sie eine Primzahl ein: __45__  
Keine Primzahl  
Geben Sie eine Primzahl ein: __12__  
Keine Primzahl  
Geben Sie eine Primzahl ein: __7__  
Primzahl  

---
Fall 2:  
Geben Sie eine Primzahl ein: __56__  
Keine Primzahl  
Geben Sie eine Primzahl ein: __zwei__  
Das ist keine Zahl.  




### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `While-Schleifen` in Python finden Sie [hier](../../50_loops/while_loop.ipynb).

<!-- grade_id: cell-ae858d4a7d4b24f8, points: 0, grade: false, solution: true -->
```python
def validiere_coupon():
    
    ###START
    
    ###ENDE
    
    return

validiere_coupon()
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-d73d157adec0fd82, points: 0, grade: true, solution: false -->
```python
try:
    validiere_coupon()
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Abfrage
Das Programm soll wie oben beschrieben funktionieren.

__Hinweis__: Diese Tests sind sehr anfällig für kleinste Abweichungen. Achten Sie bitte darauf, dass Ihr Programm *genau* so aussieht wie in der Aufgabenstellung als Beispiel dargestellt.  
Achten Sie bitte auch darauf den Text "Geben Sie eine Primzahl ein: " mittels input auszugeben (also `input("Geben Sie eine Primzahl ein: ")`). Bitte verwenden Sie hier kein print().

<!-- grade_id: cell-4384a6faf24f294d, points: 0, grade: true, solution: false -->
```python
from unittest import mock
from nose.tools import assert_equal
import unittest
import io

with mock.patch('builtins.input', side_effect=["7"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()

    print("Teste 7 als erste Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Primzahl\n', "Sollte die erste Eingabe bereits eine Primzahl sein, so sollte das Programm \"Primzahl\" ausgeben und sich beenden")
    print("Test 1 erfolgreich.")
    
with mock.patch('builtins.input', side_effect=["80", "60", "197"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()

    print("Teste 197 als dritte Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Keine Primzahl\nKeine Primzahl\nPrimzahl\n', "Das Programm soll (wenn Zahlen eingegeben werden) so lange nach einer Eingabe verlangen bis diese eine Primzahl ist. Andernfalls soll \"Keine Primzahl\" ausgegeben werden.")
    print("Test 2 erfolgreich.")
    
with mock.patch('builtins.input', side_effect=["1234", "hallo"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()

    print("Teste Buchstaben als zweite Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Keine Primzahl\nDas ist keine Zahl.\n', "Das Programm soll sich beenden wenn keine Zahl eingegeben wird. Vorher soll \"Das ist keine Zahl.\" ausgegeben werden.")
    print("Test 3 erfolgreich.")

with mock.patch('builtins.input', side_effect=["hallo"]):
    with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
        validiere_coupon()

    print("Teste Buchstaben als erste Eingabe")
    assert_equal(fake_stdout.getvalue(), 'Das ist keine Zahl.\n', "Auch wenn die erste Eingabe bereits keine Zahl ist soll das Programm die Fehlermeldung ausgeben und sich beenden.")
    print("Test 4 erfolgreich.")
```