# Aufgabe 5

Im Rahmen der Materialplanung soll ein Programm zur automatischen Berechnung mehrerer Materialbestände implementiert werden. So soll aus mehreren Listen von Eingängen und Abgängen für verschiedene Materialien jeweils der neue Materialbestand berechnet werden. 

Implementieren Sie nun die Funktion `kalkuliere_bestand`. Diese Funktion erhält als Parameter eine Liste `materialien`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `materialien` eine Liste bestehend aus dem Materialnamen als erstem Wert, gefolgt von den Eingängen und Abgängen wie gewohnt. Die Liste sieht also in etwa so aus: `[["Schrauben", +500, -200, +30], ["Nägel", +30000, -5000], ["Muttern", +20000, -3000, -500]]`.   
Als Ergebnis soll eine Liste zurückgegeben werden, auch bestehend aus mehreren Listen (`bestand`). Dabei soll jede Liste als ersten Eintrag den Materialnamen beinhalten und als zweiten Wert den Bestand dieses Materials. Für obiges Beispiel wäre das Ergebnis also `[['Schrauben', 330], ['Nägel', 25000], ['Muttern', 16500]]`.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-c4abda23c95af4d6, points: 0, grade: false, solution: true -->
```python
def kalkuliere_bestand(materialien):
    
    ###START
    
    ###ENDE
    return bestand

print("Bei \"materialien = [[\"Schrauben\", +500, -200, +30], [\"Nägel\", +30000, -5000]]\" liefert Ihre Funktion:", kalkuliere_bestand([["Schrauben", +500, -200, +30], ["Nägel", +30000, -5000]]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-947e6fe007525d14, points: 0, grade: true, solution: false -->
```python
try:
    kalkuliere_bestand([["Schrauben", +500, -200, +30], ["Nägel", +30000, -5000]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Übernahme der Materialnamen
Die Materialnamen sollen im Ergebnis jeweils als erstes Element der einzelnen Listen übernommen werden.

<!-- grade_id: cell-fc98ae5c361e6645, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

materials = [["Schrauben", +500, -200, +30], ["Nägel", +30000, -5000], ["Muttern", +20000, -500, -3000]]
stock = kalkuliere_bestand(materials)
try:
    assert_equal(stock[0][0], "Schrauben", "Der Materialname \"Schrauben\" wird nicht korrekt übernommen")
    assert_equal(stock[1][0], "Nägel", "Der Materialname \"Nägel\" wird nicht korrekt übernommen")
    assert_equal(stock[2][0], "Muttern", "Der Materialname \"Muttern\" wird nicht korrekt übernommen")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print(e)
    

materials = [["Erdbeeren", +500, -200, +30], ["Blaubeeren", +30000, -5000], ["Bananen", +20000, -500, -3000], ["Äpfel", +20]]
stock = kalkuliere_bestand(materials)
try:
    assert_equal(stock[0][0], "Erdbeeren", "Der Materialname \"Erdbeeren\" wird nicht korrekt übernommen")
    assert_equal(stock[1][0], "Blaubeeren", "Der Materialname \"Blaubeeren\" wird nicht korrekt übernommen")
    assert_equal(stock[2][0], "Bananen", "Der Materialname \"Bananen\" wird nicht korrekt übernommen")
    assert_equal(stock[3][0], "Äpfel", "Der Materialname \"Äpfel\" wird nicht korrekt übernommen")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print(e)
```

## Test 2 - Korrekte Berechnung der Materialbestände
Die Materialbestände sollen für jedes Material richtig berechnet werden.

<!-- grade_id: cell-011596bcbad4e0ad, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

materials = [["Schrauben", +500, -200, +30], ["Nägel", +30000, -5000], ["Muttern", +20000, -500, -3000]]
stock = kalkuliere_bestand(materials)
try:
    assert_equal(stock[0][1], 330, "Der Materialbestand von \"Schrauben\" wird nicht korrekt berechnet")
    assert_equal(stock[1][1], 25000, "Der Materialbestand von \"Nägel\" wird nicht korrekt berechnet")
    assert_equal(stock[2][1], 16500, "Der Materialbestand von \"Muttern\" wird nicht korrekt berechnet")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print(e)
    

materials = [["Schrauben", +700, -200, +30, +256, -13, +12], ["Nägel", +30000], ["Muttern", +20000, -500, -3000]]
stock = kalkuliere_bestand(materials)
try:
    assert_equal(stock[0][1], 785, "Der Kontostand des Kontos \"Schrauben\" wird nicht korrekt berechnet")
    assert_equal(stock[1][1], 30000, "Der Kontostand des Kontos \"Nägel\" wird nicht korrekt berechnet")
    assert_equal(stock[2][1], 16500, "Der Kontostand des Kontos \"Muttern\" wird nicht korrekt berechnet")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print(e)
```