Aufgabe 2 - Lagerbestandsberechnung
------------
Gegeben sind die Absatzprognosen sowie Produktionsmengen eines Produktes für die nächsten 12 Monate in gleichnamigen Listen.
Erstellen Sie eine neue Liste, in der Sie die Lagerbestände für die jeweiligen Monate speichern.  
Beachten Sie die Lagerbestände aus dem jeweils vorhergehenden Monat in den nächsten Monat zu übernehmen.

<!-- grade_id: cell-073db0e144d454aa, points: 0, grade: false, solution: true -->
```python
def create_stock_forecast(sales_plan, production_plan):
    
    stock_forecast = []
    previous_stock = 0
    for month in range(len(sales_plan)):
        # Berechnung des Lagerbestands für den aktuellen Monat
        previous_stock = stock_forecast[month]
        
    return stock_forecast

print("Bei \"sales_plan = [20, 30, 40]\" und \"production_plan = [50, 60, 65]\" liefert Ihre Funktion:", create_stock_forecast([20, 30, 40], [50, 60, 65]))
```

   ***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-264810029e59bd40, points: 0, grade: true, solution: false -->
```python
try:
    create_stock_forecast([0, 0, 0, 0], [0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Länge der Ergebnisliste
Das von der Funktion `create_stock_forecast` zurückgegebene Ergebnis  `stock_forecast` ist eine Liste.
Diese Liste muss genau so viele Einträge enthalten wie
die übergebenen Listen `sales_plan` und `production_plan`. In diesem Test hat die übergebene Liste die Länge 12.

<!-- grade_id: cell-b1d2a40bbb45242f, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

sales_plan = [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
production_plan = [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
stock_forecast = create_stock_forecast(sales_plan, production_plan)
try:
    assert_is_instance(stock_forecast, list)
    assert_equal(len(stock_forecast), 12)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Die Lagerbestände sollen als Liste ausgegeben werden.\nDie zurückgegebene Liste muss 12 Einträge enthalten.")
    print(e)
```

Output:
```
Test 1 erfolgreich
```

## Test 2 - Korrekte Berechnung der Lagerbestände
Die von der Funktion `create_stock_forecast` zurückgegebene Liste `stock_forecast` muss die korrekten Lagerbestände in
Abhängigkeit von den übergebenen Absatzmengen und Produktionsmengen enthalten.

<!-- grade_id: cell-504e093f2573e140, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_is_instance

sales_plan = [570, 500, 520, 500, 600, 550, 530, 520, 530, 590, 600, 590]
production_plan = [600, 520, 520, 520, 650, 650, 500, 530, 500, 600, 620, 610]
try:
    expected_stock_forecast = [30, 50, 50, 70, 120, 220, 190, 200, 170, 180, 200, 220]
    assert_equal(create_stock_forecast(sales_plan, production_plan), expected_stock_forecast)
    print("Test 2 erfolgreich")
except AssertionError as e:
print("Die berechneten Lagerbestände sind nicht korrekt.")
print(e)
```

Output:
```
Test 2 erfolgreich
```