# Aufgabe 3

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Analyse der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden. So soll aus einer Liste mit vielen Einnahmen heruasgefiltert werden, welches die höchste eingenommene Umsatzsteuer ist. Außerdem sollen der durchschnitt sowie die Summe aller eingenommenen Umsatzsteuern berechnet werden.

Implementieren Sie nun die Funktion `analyse_taxes`. Diese Funktion erhält als Parameter eine Liste `sales`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).   
Als Ergebnis soll der gesamte Betrag, welcher durch die Umsatzsteuer eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden ("taxes").   
Die durchschnittliche Höhe der Umsatzsteuer ("average").  
Der höchste Betrag, welcher durch die Umsatzsteuer an einem Artikel eingenommen wurde, in Form eines Float-Wertes ("highest_tax").   

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-ea4effec3a83d25e, points: 0, grade: false, solution: true -->
```python
def analyse_taxes(sales):
    
    ###BEGIN SOLUTION
    taxes = 0
    highest_tax = 0
    for sale in sales:
        # Berechne mit jedem Durchlauf eine neue Summe für "taxes"
        # Prüfe außerdem ob die einzelne Umsatzsteuer die höchste ist
            
    average = taxes / len(sales)
    
    ###END SOLUTION
    return taxes, average, highest_tax

print("Bei \"sales = [20, 30, 40]\" liefert Ihre Funktion:", analyse_taxes([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-8baa31006b25d951, points: 0, grade: true, solution: false -->
```python
try:
    analyse_taxes([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung der gesamten Umsatzsteuer
Die Umsatzsteuer (`taxes`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-6a7e886faa4b80e2, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300, 250.45, 19.78, 123.56, 50.23]
taxes = analyse_taxes(sales)
try:
    assert_equal(taxes[0], 141.3638)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)

    
sales = [99999999, 1234567, 9876543]
taxes = analyse_taxes(sales)
try:
    assert_equal(taxes[0], 21111110.71)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)

    
sales = [12, 20]
taxes = analyse_taxes(sales)
try:
    assert_equal(taxes[0], 6.08)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Berechnung des Durchschnitts
Der Durchschnitt (`average`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-70e120823526c483, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [500, 250.45, 19.78, 123.56, 50.23]
average = analyse_taxes(sales)
try:
    assert_equal(average[1], 35.87276, "Das Ergebnis ist falsch")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [99999999, 1234567, 9876543]
average = analyse_taxes(sales)
try:
    assert_equal(average[1], 7037036.903333333, "Das Ergebnis ist falsch")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [12, 20]
average = analyse_taxes(sales)
try:
    assert_equal(average[1], 3.04)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [20]
average = analyse_taxes(sales)
try:
    assert_equal(average[1], 3.8)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```

## Test 3 - Korrekte Bestimmung der höchsten Umsatzsteuer
Der Betrag (`highest_tax`) soll für jede Liste korrekt herausgefiltert werden.

<!-- grade_id: cell-bd0f620abcd1ffda, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300, 250.45, 19.78, 123.56, 50.23]
highest_tax = analyse_taxes(sales)
try:
    assert_equal(highest_tax[2], 57)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [99999999, 1234567, 9876543]
highest_tax = analyse_taxes(sales)
try:
    assert_equal(highest_tax[2], 18999999.81)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [12, 20]
highest_tax = analyse_taxes(sales)
try:
    assert_equal(highest_tax[2], 3.8)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

sales = [20]
highest_tax = analyse_taxes(sales)
try:
    assert_equal(highest_tax[2], 3.8)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```