# Aufgabe 2

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Analyse von Materialeingängen und Materialabgängen implementiert werden. So soll aus einer großen Liste von Eingängen und Abgängen erkannt werden welche Transaktion die größte ist.  

Implementieren Sie nun die Funktion `analyse_bestand`. Diese Funktion erhält als Parameter eine Liste `eingehend_ausgehend`, welche in Form von Integer-Werten die Eingänge und Abgänge eines Materials in eine Lagerhalle angibt (in etwa `[+500, -200, +70, -65]`).   
Als Ergebnis soll die höchste Transaktion(`hoechste_transaktion`) als Integer-Wert zurückgegeben werden. Hierbei sollen sowohl Eingänge als auch Abgänge berücksichtigt werden.



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-ac48f4aedf1f6935, points: 0, grade: false, solution: true -->
```python
def analyse_bestand(eingehend_ausgehend):
    
    ###START
    
    ###ENDE
    return hoechste_transaktion

print("Bei \"eingehend_ausgehend = [20, 30, 40]\" liefert Ihre Funktion:", analyse_bestand([20, 30, 40]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-bfa0e7649eab31c7, points: 0, grade: true, solution: false -->
```python
try:
    analyse_bestand([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bestimmung der höchsten Materialtransaktion
Die höchste Warenbewegung (`hoechste_transaktion`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch Materialabgänge zu berücksichtigen.

<!-- grade_id: cell-b2c5e9ad80aa9392, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+100, -26, -23, +210, +234]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction, 234)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+200, -200, +200, -200]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction, 200)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+15]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction, 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+300, -500, -23, +210, +234]
highest_transaction = analyse_bestand(incoming_outgoing)
try:
    assert_equal(highest_transaction, -500)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```