# Aufgabe 5

Es soll ein Programm zur automatischen Berechnung mehrerer Kontostände implementiert werden. So soll aus mehreren Listen von Zugängen und Abgängen für verschiedene Konten jeweils der neue Kontostand berechnet werden. 

Implementieren Sie nun die Funktion `berechne_kontostaende`. Diese Funktion erhält als Parameter eine Liste `ueberweisungen`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `ueberweisungen` eine Liste bestehend aus dem Kontonamen als erstem Wert, gefolgt von den Zugängen und Abgängen wie gewohnt. Die Liste sieht also in etwa so aus: `[["Kasse", +500, -200, +30], ["Immobilien", +30000, -5000], ["Fuhrpark", +500, -20000, +3000]]`.   
Als Ergebnis soll eine Liste zurückgegeben werden, auch bestehend aus mehreren Listen. Dabei soll jede Liste als ersten Eintrag den Kontonamen beinhalten und als zweiten Wert den Kontostand dieses Kontos. Für obiges Beispiel wäre das Ergebnis also `[['Kasse', 330], ['Immobilien', 25000], ['Fuhrpark', -16500]]`.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-7ecc30079e6674c8, points: 0, grade: false, solution: true -->
```python
def berechne_kontostaende(ueberweisungen):
    
    ###START
    
    ###END SOLUTION
    return kontostaende

print("Bei \"ueberweisungen = [[\"Kasse\", +500, -200, +30], [\"Immobilien\", +30000, -5000]]\" liefert Ihre Funktion:", berechne_kontostaende([["Kasse", +500, -200, +30], ["Immobilien", +30000, -5000]]))
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-65aa12cf9dce6952, points: 0, grade: true, solution: false -->
```python
try:
    berechne_kontostaende([["Kasse", +500, -200, +30], ["Immobilien", +30000, -5000]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Übernahme der Kontennamen
Die Kontennamen sollen im Ergebnis jeweils als erstes Element der einzelnen Listen übernommen werden.

<!-- grade_id: cell-9f31a47ff04a715b, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [["Kasse", +500, -200, +30], ["Immobilien", +30000, -5000], ["Fuhrpark", +500, -20000, +3000]]
balances = berechne_kontostaende(transfers)
try:
    assert_equal(balances[0][0], "Kasse", "Der Kontoname \"Kasse\" wird nicht korrekt übernommen")
    assert_equal(balances[1][0], "Immobilien", "Der Kontoname \"Immobilien\" wird nicht korrekt übernommen")
    assert_equal(balances[2][0], "Fuhrpark", "Der Kontoname \"Fuhrpark\" wird nicht korrekt übernommen")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print(e)
    

transfers = [["Verbindlichkeiten", +500, -200, +30], ["Forderungen", +30000, -5000], ["Bürobedarf", +500, -20000, +3000], ["Lager", +200, -200, +200, -200]]
balances = berechne_kontostaende(transfers)
try:
    assert_equal(balances[0][0], "Verbindlichkeiten", "Der Kontoname \"Verbindlichkeiten\" wird nicht korrekt übernommen")
    assert_equal(balances[1][0], "Forderungen", "Der Kontoname \"Forderungen\" wird nicht korrekt übernommen")
    assert_equal(balances[2][0], "Bürobedarf", "Der Kontoname \"Bürobedarf\" wird nicht korrekt übernommen")
    assert_equal(balances[3][0], "Lager", "Der Kontoname \"Lager\" wird nicht korrekt übernommen")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print(e)
```

## Test 2 - Korrekte Berechnung der Kontostände
Die Kontostände sollen für jedes Konto richtig berechnet werden.

<!-- grade_id: cell-5fa4a39ed84279fb, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [["Kasse", +500, -200, +30], ["Immobilien", +30000, -5000], ["Fuhrpark", +500, -20000, +3000]]
balances = berechne_kontostaende(transfers)
try:
    assert_equal(balances[0][1], 330, "Der Kontostand des Kontos \"Kasse\" wird nicht korrekt berechnet")
    assert_equal(balances[1][1], 25000, "Der Kontostand des Kontos \"Immobilien\" wird nicht korrekt berechnet")
    assert_equal(balances[2][1], -16500, "Der Kontostand des Kontos \"Fuhrpark\" wird nicht korrekt berechnet")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print(e)
    

transfers = [["Kasse", +700, -200, +30, +256, -13, +12], ["Immobilien", +30000], ["Fuhrpark", +500, -20000, +3000]]
balances = berechne_kontostaende(transfers)
try:
    assert_equal(balances[0][1], 785, "Der Kontostand des Kontos \"Kasse\" wird nicht korrekt berechnet")
    assert_equal(balances[1][1], 30000, "Der Kontostand des Kontos \"Immobilien\" wird nicht korrekt berechnet")
    assert_equal(balances[2][1], -16500, "Der Kontostand des Kontos \"Fuhrpark\" wird nicht korrekt berechnet")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print(e)
```