# Aufgabe 4

Im Rahmen der operativen Planung in einem Handelsunternehmen soll auf der Basis unterschiedlicher Faktoren berechent werden, ob eine Handelsware vom Lieferanten beschafft werden muss.

ie Beschaffung wird in dieser Aufgabe über einen Meldebestand sowie manuelle Bestellanforderungen gesteuert. Hierfür sollen folgende Bedingungen gelten:

- Ist der Bestand einer Handelsware kleiner oder gleich dem Meldebestand, werden 1000 Stück der Handelsware bestellt
- Falls eine manuelle Bestellanforderungen existeirt werden 500 Stück der Handelsware bestellt.  
 
In beiden Fällen soll zusätzlich geprüft werden ob durch den Wert der Bestellung der Kreditrahmen beim Lieferanten überschritten würde. Falls der Kreditrahmen überschritten wird soll die Bestellmenge 0 sein. **Nutzen Sie zur Lösung 
der Aufgabe komplexe Bedingungen. Verwenden Sie keine verschachtelten ìf`-Anweisungen.**


Implementieren Sie die Funktion `bestellmenge_berechnen`. Die Funktion erhält folgende Parameter:
- `meldebestand`: Der Paramter enthält den Meldebestand für eine Handelsware und ist vom Typ `Integer`.
- `bestand`: Der Parameter enthält den aktuellen Bestand und ist vom Typ `Integer`.
- `manuelle_bestellanforderung`: Der Parameter ist vom Typ `Boolean` und enthält den Wert `True` falls eine maunelle Bestellanforderung existiert.
- `kreditrahmen`: Der Parameter enthält den aktuellen Kreditrahme bei Lieferanten.
- `preis`: Der Parameter enthält den Preis der Handelsware und ist vom Typ `Float`.

Als Ergebnis der Funktion `bestellmenge_berechnen` sollen die Variable `bestellmenge` berechnet werden. Die Variable `bestellmenge` gibt an wie viele Stücke einer Handelsware beschafft werden sollen.

## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

## Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

```python
def bestellmenge_berechnen(meldebestand, bestand, manuelle_bestellanfoderung, kreditrahmen, preis):
    
    ###BEGIN SOLUTION
    ### START
    bestellmenge = 0

    if bestand <= meldebestand and kreditrahmen >= 1000 * preis:
        bestellmenge = 1000
    elif manuelle_bestellanfoderun and kreditrahmen >= 500 * preis:
        bestellmenge = 500

    ### ENDE
    ###END SOLUTION
    return bestellmenge

print("Wenn Ihre Funktion mit den Parametern \"meldebestand = 500\", \"bestand = 499\", \"anuelle_bestellanfoderung = False\", \"kreditrahmen = 99999\" und \"preis = 2.25\" aufgerufen wird ergibt sich eine Bestellmenge von", bestellmenge_berechnen(500, 499, False, 99999, 2.25))
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


## Test 1 - Korrekte Bedingungen bei ausreichendem Kreditrahmen
Die Menge an zu bestellenden Produkten (`bestellmenge`) soll 1000 sein wenn der Meldebestand erreicht ist.  
Die Menge soll 500 sein wenn eine manueller Bestellanforderung vorliegt.  
In beiden Fällen ist die Voraussetzung, dass der Kretitrahmen ausreichend ist. 

```python
from nose.tools import assert_equal
from utils import *

credit = 999999
reorder_level = 500
stock_today = 499
manual_reorder = False
try:
    reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)

    assert_equal(reorder_quantity, 1000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde")
    print(e)
except Exception as e:
    print(friendly_exceptions.msg(e)) 
    
reorder_level = 500
stock_today = 501
manaul_reorder = False
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn der Meldebestand nicht erreicht wurde und keine manueller Bestellwunsch geäußert wurde")
    print(e)
    
    
reorder_level = 500
stock_today = 501
manaul_reorder = True
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 500 sein wenn ein manueller Bestellwunsch vorliegt aber der Meldebestand nicht erreicht wurde.")
    print(e)
    

reorder_level = 500
stock_today = 499
manaul_reorder = True
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 1000)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde, unabhängig davon ob ein Bestellwunsch vorliegt")
    print(e)
```

## Test 2 - Korrekte Bedingungen mit begrenztem Guthaben
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll falls das Guthaben nicht reicht 0 sein.  
Sollte das Guthaben nur für die große Bestellung nicht reichen aber eine manuelle Bestellung vorliegen, so sollen 500 Stück bestellt werden.

<!-- grade_id: cell-6c6e3aadbcf7b04d, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

credit = 3000
reorder_level = 500
stock_today = 499
manual_reorder = False
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 1000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde und genug Guthaben vorhanden ist")
    print(e)
    
    
credit = 500
reorder_level = 500
stock_today = 499
manual_reorder = False
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn der Meldebestand erreicht wurde aber zu wenig Guthaben vorhanden ist")
    print(e)
    

credit = 1500
reorder_level = 500
stock_today = 499
manual_reorder = True
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder, credit)
try:
    assert_equal(reorder_quantity, 500)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 500 sein wenn das Guthaben zwar nicht für 1000Stück reicht aber für 500 und eine manuelle Bestellung vorliegt")
```