# Aufgabe 3

Im Rahmen der operativen Planung in einem Handelsunternehmen soll erneut auf der Basis mehrerer gegebener
Faktoren eine Einschätzung bezüglich etwaiger Nachbestellungen für ein Produkt getroffen werden. 

Ist die Lagermenge eines bestimmten Produktes unter oder gleich dem Meldebestand, so soll dieses Produkt automatisch 1000 mal nachbestellt werden. Allerdings, sollte das Produkt nicht den Meldebestand erreicht haben, so muss geprüft werden ob ein manuell eingetragener Wunsch zur Nachbestellung vorliegt. Wenn ja wird halb so viel wie normalerweise bestellt (500). Sollte keiner der beiden Fälle eintreten soll außerdem "Status OK" mittels print() in die Konsole ausgegeben werden.

Implementieren Sie nun die Funktion `reorder`. Die Funktion erhält als Parameter `reorder_level` einen Integer Wert, welcher den Meldebestand angibt.  
Sowie `stock_today`, welcher angibt wie oft das Produkt an diesem Morgen noch im Lager vorrätig ist.  
Der dritte Parameter `manual_reorder` ist vom Typ Boolean und ist sobald ein manuel eingetragener Nachbestellwunsch vorliegt = True.  
Als Ergebnis soll eine neuer Integer Wert `reorder_quantity` erstellt und zurückgegeben werden, welcher angibt wie viele Exemplare des Produktes nachbestellt werden sollen. Soll es nicht nachbestellt werden, so ist die `reorder_quantity` = 0.

## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-34146920188d4f1e, points: 0, grade: false, solution: true -->
```python
def reorder(reorder_level, stock_today, manual_reorder):
    ###START
    
    reorder_quantity = 0
    if stock_today <= reorder_level:
        reorder_quantity = 1000
    elif #Bedingung:
        reorder_quantity = 500
    else:
        print("Status OK")
    
    ###ENDE
    return reorder_quantity


print("Wenn Ihre Funktion mit den Parametern \"reorder_level = 500\", \"stock_today = 499\", \"manual_reorder = False\" aufgerufen wird ist das Ergebnis:", reorder(500, 499, False))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "reorder_level = 500", "stock_today = 499", "manual_reorder = False" aufgerufen wird ist das Ergebnis: 1000
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-14d1e232dfce8578, points: 0, grade: true, solution: false -->
```python
try:
    reorder(500, 499, False)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingungen
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll 1000 sein wenn der Meldebestand erreicht ist.  
Es soll außerdem 500 sein wenn ein manueller Bestellwunsch vorliegt oder 0 wenn keiner vorliegt und der Meldebestand nicht erreicht ist.

<!-- grade_id: cell-89c9f525adde2701, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

reorder_level = 500
stock_today = 499
manual_reorder = False
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder)
try:
    assert_equal(reorder_quantity, 1000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde")
    print(e)
    
    
reorder_level = 500
stock_today = 501
manaul_reorder = False
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn der Meldebestand nicht erreicht wurde und keine manueller Bestellwunsch geäußert wurde")
    print(e)

    
reorder_level = 500
stock_today = 501
manaul_reorder = True
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 500 sein wenn ein manueller Bestellwunsch vorliegt aber der Meldebestand nicht erreicht wurde.")
    print(e)
    
    
reorder_level = 500
stock_today = 499
manaul_reorder = True
reorder_quantity = reorder(reorder_level, stock_today, manual_reorder)
try:
    assert_equal(reorder_quantity, 1000)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde, unabhängig davon ob ein Bestellwunsch vorliegt")
    print(e)
```

Output:
```
Test 1 erfolgreich
Status OK
Test 2 erfolgreich
Status OK
Test 3 erfolgreich
Test 4 erfolgreich
```