# Aufgabe 4 

In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Im vierten Schritt soll weiter zwischen den Risikoklassen differenziert werden. Ein Gebäude soll immer dann in Risikoklasse "D" eingestuft werden wenn es sich um ein Mehrfamilienhaus handelt *oder* es älter als 20 Jahre ist. Andernfalls soll folgendes überprüft werden:  
Ist das Alter eines Gebäudes unter oder gleich 10 Jahren, so gehört es in die Risikoklasse "A".  
Ist ein Gebäude im Alter von 11-20 Jahren, so ist die Risikoklasse "B". 

Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren). Sowie den Parameter `is_apartement_building`, der in Form eines Boolean-Wertes angibt ob das Gebäude ein Mehrfamilienhaus ist (True oder False).  
Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden ("A","B" oder "D"). 


## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-193d12ab4b908fe7, points: 0, grade: false, solution: true -->
```python
def insurance_classification(age, is_apartement_building):
    
    ###START
    insurance_class = "A"
    if #Bedingung, nach der das Gebäude in Klasse D gestuft wird:
        insurance_class = "D"
    elif #Bedingung nach der das Gebäude in Klasse A gestuft wird:
        insurance_class = "A"
    else:
        #Einstufung in andere Klasse
        
    ###ENDE
    return insurance_class


print("Wenn Ihre Funktion mit den Parametern \"age = 5\" und \"is_apartement_building = True\" aufgerufen wird ist das Ergebnis:", insurance_classification(5, True))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "age = 5" und "is_apartement_building = True" aufgerufen wird ist das Ergebnis: D
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-b3a0c892c746e20b, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(5, True)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingungen
Die Risikoklasse (`insurance_class`) soll "D" sein, wenn es sich bei dem Gebäude um ein Mehrfamilienhaus handelt *oder* das Gebäude älter als 20 Jahre ist.  
Ist dies nicht der Fall so soll die Risikoklasse "B" oder "A" sein, je nachdem ob das Gebäude jünger als 10 Jahre ist (A) oder zwischen 11 und 20 Jahren (B).

<!-- grade_id: cell-4002fd512e8fb9f7, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 5
is_apartement_building = True
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
    print(e)
    
    
age = 25
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude älter als 20 Jahre ist.")
    print(e)
    

age = 15
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "B")
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"B\" sein wenn das Gebäude kein Mehrfamilienhaus ist und zwischen 11 und 20 Jahren alt.")
    print(e)
    

age = 5
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude kein Mehrfamilienhaus ist und noch nicht älter als 10 Jahre.")
    print(e)
    
    
age = 25
is_apartement_building = True
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 5 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude sowohl ein Mehrfamilienhaus, als auch älter als 20 ist.")
    print(e)

```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
Test 5 erfolgreich
```

## Test 2 - Beachtung der Randfälle
Die Risikoklasse (`insurance_class`) soll auch dann "A" sein wenn das Gebäude kein Mehrfamilienhaus ist und *genau* 10 Jahre alt ist. Ebenso gehört ein *genau* 20 Jahre altes Gebäude noch zur Risikoklasse "B".

<!-- grade_id: cell-3df2e0a21b94f1cc, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 10
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein auch wenn das Gebäude genau 10 Jahre alt ist.")
    print(e)
    
    
age = 20
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "B")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"B\" sein auch wenn das Gebäude genau 20 Jahre alt ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```