# Aufgabe 2

In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Im zweiten Schritt wollen Sie überprüfen ob das Alter eines Gebäudes unter oder gleich 10 Jahren ist. Ist dies der Fall so ist die Risikoklasse in jedem Fall "A".  

Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren).  
Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden. Ist das Gebäude älter als 10 Jahre, so geben Sie "D" zurück. Ansonsten geben Sie "A" zurück. Zusätzlich geben Sie bitte mittels print() in die Konsole aus ob das Gebäude als "neuwertig" oder "alt" eingestuft wurde.


## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-3d302f7daaba320e, points: 0, grade: false, solution: true -->
```python
def insurance_classification(age):
    
    insurance_class = "D"
    if #Bedingung (alter unter oder gleich 10):
        #Risikoklasse setzen
        print("Eingestuft als neuwertig.")
    else:
        #Risikoklasse setzen
        print("Eingestuft als alt.")
    
    return insurance_class


print("Wenn Ihre Funktion mit dem Parameter \"age = 5\" aufgerufen wird ist das Ergebnis:", insurance_classification(5))
```

Output:
```
Eingestuft als neuwertig.
Wenn Ihre Funktion mit dem Parameter "age = 5" aufgerufen wird ist das Ergebnis: A
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-3210983ce2669555, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(5)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

Output:
```
Eingestuft als neuwertig.
```

## Test 1 - Korrekte Bedingung
Die Risikoklasse (`insurance_class`) soll "D" sein wenn das Gebäude älter als 10 Jahre ist. Ansonsten soll sie "A" sein.

<!-- grade_id: cell-321179758e62c081, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 20
insurance_class = insurance_classification(age)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude älter als 10 Jahre ist.")
    print(e)
    
    
age = 5
insurance_class = insurance_classification(age)
try:
    assert_equal(insurance_class, "A")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude noch keine 10 Jahre alst ist.")
    print(e)

```

Output:
```
Eingestuft als alt.
Test 1 erfolgreich
Eingestuft als neuwertig.
Test 2 erfolgreich
```

## Test 2 - Beachtung der Randfälle
Die Risikoklasse (`insurance_class`) soll auch dann "A" sein wenn das Gebäude *genau* 10 Jahre alt ist. Jedoch ab 11 Jahren soll die Klasse "D" sein.

<!-- grade_id: cell-709ebc1a87ed56ae, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 11
insurance_class = insurance_classification(age)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein auch wenn das Gebäude nur knapp älter als 10 Jahre ist.")
    print(e)
    
    
age = 10
insurance_class = insurance_classification(age)
try:
    assert_equal(insurance_class, "A")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein auch wenn das Gebäude genau 10 Jahre alst ist.")
    print(e)
```

Output:
```
Eingestuft als alt.
Test 1 erfolgreich
Eingestuft als neuwertig.
Test 2 erfolgreich
```