# Aufgabe 1 

Im Rahmen der operativen Planung in einem Handelsunternehmen soll auf der Basis mehrerer gegebener
Faktoren eine Einschätzung bezüglich etwaiger Nachbestellungen für ein Produkt getroffen werden. 

Wurde ein Produkt am Vortag mindestens 100 mal verkauft, so soll dieses gleich 2-fach so oft nachbestellt werden wie es am Vortag verkauft wurde.

Implementieren Sie nun die Funktion `reorder`. Die Funktion erhält als Parameter `sales_yesterday` einen Integer Wert, welcher angibt wie oft ein Produkt am Vortag verkauft wurde. Als Ergebnis soll eine neuer Integer Wert `reorder_quantity` erstellt und zurückgegeben werden, welcher angibt wie viele Exemplare des Produktes nachbestellt werden sollen. Soll es nicht nachbestellt werden, so ist die `reorder_quantity` = 0.

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-4b9fd4e745d85dde, points: 0, grade: true, solution: true -->
```python
def reorder(sales_yesterday):
    
    ###BEGIN SOLUTION
    reorder_quantity = 0
    if sales_yesterday >= 100:
        reorder_quantity = sales_yesterday * 2
    ###END SOLUTION
    return reorder_quantity

print("Wenn Ihre Funktion mit dem Parameter \"sales_yesterday = 500\" aufgerufen wird ist das Ergebnis:", reorder(500))
```

Output:
```
Wenn Ihre Funktion mit dem Parameter "sales_yesterday = 500" aufgerufen wird ist das Ergebnis: 1000
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-b5a615ac2b243b94, points: 0, grade: true, solution: false -->
```python
try:
    reorder(500)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingung
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll sich ändern wenn das Produkt am Vortag mindestens 100 mal verkauft wurde. Ansonsten soll sie gleich 0 sein.

<!-- grade_id: cell-bc1532af7f2c9581, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

sales_yesterday = 50
reorder_quantity = reorder(sales_yesterday)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn das Produkt weniger als 100 mal verkauft wurde.")
    print(e)
    
    
sales_yesterday = 500
reorder_quantity = reorder(sales_yesterday)
try:
    assert_not_equal(reorder_quantity, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll nicht 0 sein wenn das Produkt nachbestellt werden muss.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```

## Test 2 - Beachtung der Randfälle
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll sich ändern wenn das Produkt am Vortag **mindestens** 100 mal verkauft wurde.  
Das heißt sollte es **genau** 100 mal verkauft worden sein, so soll auch eine Bestellmenge != 0 angegeben werden.

<!-- grade_id: cell-9f37475cda6603bd, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

sales_yesterday = 100
reorder_quantity = reorder(sales_yesterday)
try:
    assert_not_equal(reorder_quantity, 0)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll nicht 0 sein auch wenn die Bedingung nur gerade so erfüllt ist")
    print(e)
```

Output:
```
Test 1 erfolgreich
```

## Test 3 - Korrekte Berechnung der Bestellmenge
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll korrekt berechnet werden. Nämlich das 2-fache der Verkaufsmenge des Vortages

<!-- grade_id: cell-b7ba8a17f5efa24a, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

sales_yesterday = 345
reorder_quantity = reorder(sales_yesterday)
try:
    assert_equal(reorder_quantity, 690)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 2 mal so groß sein wie \"sales_yesterday\".")
    print(e)
```

Output:
```
Test 1 erfolgreich
```