 # Aufgabe 2

 In einer Bank sollen Kreditkartenanträge zukünftig mittels eines Computerprogramms automatisch bearbeitet werden. Sie sollen nun das Regelwerk zur Entscheidung über Anträge mittels `Conditionals` in einem Python Programm umsetzen.
 Im zweiten Schritt soll überprüft werden ob das Einkommen eines Antragstellers über 1500€ liegt. Nur dann soll eine Kreditkarte ausgestellt werden. Andernfalls ist der Antragsteller nicht berechtigt eine Kreditkarte zu erhalten.

 Implementieren Sie nun die Funktion `request_verification`. Diese Funktion erhält einen Parameter `monthly_income`, welcher in Form eines float-Wertes das monatliche Einkommen des Antragstellers speichert.
 Als Ergebnis soll in Form eines Boolean-Wertes zurückgegeben werden ob der Antragsteller berechtigt ist eine Kreditkarte zu erhalten. Ist das monatliche Einkommen unter oder gleich 1500€ so geben Sie `False`zurück. Ansonsten `True`.
 **Geben Sie zusätzlich mittels print() eine Meldung in die Konsole aus ob der Antrag angenommen oder abgelehnt wird.**



 ### Hilfe
 Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
 Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-435d35e71c19e9e4, points: 0, grade: false, solution: true -->
```python
def request_verification(monthly_income):
    
    is_eligible = True
    if #Bedingung (Einkommen unter oder gleich 1500€):
        is_eligible = False
        print("Antrag abgelehnt")
    else:
        is_eligible = True
        print("Antrag angenommen")
        
    return is_eligible


print("Wenn Ihre Funktion mit dem Parameter \"monthly_income = 500\" aufgerufen wird ist das Ergebnis:", request_verification(500))
```

Output:
```
Antrag abgelehnt
Wenn Ihre Funktion mit dem Parameter "monthly_income = 500" aufgerufen wird ist das Ergebnis: False
```

 ***
 # Automatische Tests
 Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-985a5692173d952d, points: 0, grade: true, solution: false -->
```python
try:
    request_verification(500)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

Output:
```
Antrag abgelehnt
```

 ## Test 1 - Rückgabe als Boolean-Wert
 Der Rückgabewert (`is_eligible`) soll vom Datenypen Boolean sein.

<!-- grade_id: cell-13874ea7e8a804fa, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 3000
is_eligible = request_verification(monthly_income)
try:
    assert_equal(isinstance(is_eligible, bool), True)
    assert_equal(isinstance(is_eligible, str), False)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll vom Typen Boolean sein.\nDas Ergebnis soll nicht vom Typen String sein. Nutzen Sie keine Anführungszeichen für die Zuweisung von Boolean-Werten")
    print(e)
    
    
```

Output:
```
Antrag angenommen
Test 1 erfolgreich
```

 ## Test 2 - Korrekte Bedingung
 Der Rückgabewert (`is_eligible`) soll *True* sein wenn der Antragsteller ein monatliches Einkommen von über 1500€ hat. Ansonsten soll der Wert *False* sein.

<!-- grade_id: cell-10bc2b025630b9e3, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 3000
is_eligible = request_verification(monthly_income)
try:
    assert_equal(is_eligible, True)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll True sein wenn der Antragsteller ein Einkommen von über 1500€ hat.")
    print(e)
    
    
monthly_income = 200
is_eligible = request_verification(monthly_income)
try:
    assert_equal(is_eligible, False)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll False sein wenn der Antragsteller ein Einkommen von unter 1500€ hat.")
    print(e)
    
    
```

Output:
```
Antrag angenommen
Test 1 erfolgreich
Antrag abgelehnt
Test 2 erfolgreich
```

 ## Test 3 - Beachtung von Randfällen
 Der Rückgabewert (`is_eligible`) soll auch dann *False* sein wenn der Antragsteller ein monatliches Einkommen von genau 1500€ hat. Jedoch soll der Wert schon *True* sein wenn das Einkommen nur minimal über 1500€ liegt.

<!-- grade_id: cell-3883c4ef269a897c, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 1500
is_eligible = request_verification(monthly_income)
try:
    assert_equal(is_eligible, False)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann False sein wenn der Antragsteller ein Einkommen von genau 1500€ hat.")
    print(e)
    
    
monthly_income = 1500.01
is_eligible = request_verification(monthly_income)
try:
    assert_equal(is_eligible, True)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann True sein wenn der Antragsteller ein Einkommen von nur knapp über 1500€ hat.")
    print(e)
    
    
```

Output:
```
Antrag abgelehnt
Test 1 erfolgreich
Antrag angenommen
Test 2 erfolgreich
```