# Aufgabe 5 

In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Setzen Sie in diesem Schritt folgendes Regelwerk um:
- Falls das Alter eines Hauses kleiner oder gleich 10 Jahre ist und es sich um ein Reihenhaus oder Einfamilienhaus   handelt, ist die Risikoklasse A 
- Falls es sich um ein Reihenhaus in Alter von 11 - 20 Jahren handelt ist die Risikoklasse B 
- Falls es sich um ein Einfamilienhaus im Alter zwischen 11 – 15 Jahre handelt ist die Risikoklasse B 
- Falls es sich um ein Einfamilienhaus im Alter zwischen 16 – 20 Jahre handelt ist die Risikoklasse C 
- Bei Mehrfamilienhäusern ist die Risikoklasse D 
- Bei Häusern ab einem Alter von 21 Jahren ist die Risikoklasse D. 

Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren). Außerdem den Parameter `building_type`, der in Form eines Strings angibt ob das Gebäude ein "detached_house" (Einfamilienhaus), "townhouse" (Reihenhaus) oder "apartement_building" (Mehrfamilienhaus) ist.  
**Testen Sie bevor Sie das Regelwerk umsetzen ob der Parameter "building_type" tatsächlich ein String ist damit Sie diesen verwenden können. Andernfalls geben Sie als Risikoklasse "ERROR" zurück.**  
Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden ("A","B", "C" oder "D"). 


## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-44dd4fd08f2e3fce, points: 0, grade: false, solution: true -->
```python
def insurance_classification(age, building_type):
    
    ###START
    insurance_class = "ERROR"
    if isinstance(building_type, str):
        if building_type == "apartement_building" or #weitere Bedingung für Risikoklasse D (alter über 20):
            insurance_class = "D"
        elif age <= 10:
            insurance_class = "A"
        elif building_type == "townhouse" or #weitere Bedingung für Risikoklasse B (alter zwischen 11 und 15):
            insurance_class = "B"
        else:
            insurance_class = "C"
        
        
    ###ENDE
    return insurance_class


print("Wenn Ihre Funktion mit den Parametern \"age = 5\" und \"building_type = townhouse\" aufgerufen wird ist das Ergebnis:", insurance_classification(5, "townhouse"))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "age = 5" und "building_type = townhouse" aufgerufen wird ist das Ergebnis: A
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-fd946f16b75ae09c, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(5, "townhouse")
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - ERROR bei fehlerhaftem Parameter
Sollte der Parameter `building_type` kein String sein, so muss dies als "ERROR" von Ihrem Programm abgefangen werden. Sonst ist der weiter Umgang mit dem Parameter nicht möglich.

<!-- grade_id: cell-ea31930496ec0f0c, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 5
building_type = 15
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "ERROR")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"ERROR\" sein wenn der Parameter `building_type` kein String ist.")
    print(e)
    
    
age = 5
building_type = True
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "ERROR")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"ERROR\" sein wenn der Parameter `building_type` kein String ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```

## Test 2 - Korrekte Bedingungen

<!-- grade_id: cell-6ab3b525d6f12ab1, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 5
building_type = "townhouse"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "A")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude ein Reihenhaus, jünger als 10 Jahre ist.")
    print(e)
    
    
age = 5
building_type = "detached_house"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "A")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude ein Einfamilienhaus, jünger als 10 Jahre ist.")
    print(e)
    
    
age = 15
building_type = "townhouse"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "B")
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"B\" sein wenn das Gebäude ein Reihenhaus, zwischen 11-20 Jahren alt ist.")
    print(e)
    
    
age = 13
building_type = "detached_house"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "B")
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"B\" sein wenn das Gebäude ein Einfamilienhaus, zwischen 11-15 Jahren alt ist.")
    print(e)
    

age = 17
building_type = "detached_house"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "C")
    print("Test 5 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"C\" sein wenn das Gebäude ein Einfamilienhaus, zwischen 16-20 Jahren alt ist.")
    print(e)

    
age = 0
building_type = "apartement_building"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "D")
    print("Test 6 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
    print(e)
    
    
age = 25
building_type = "townhouse"
insurance_class = insurance_classification(age, building_type)

try:
    assert_equal(insurance_class, "D")
    print("Test 7 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude 21 Jahre oder älter ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
Test 5 erfolgreich
Test 6 erfolgreich
Test 7 erfolgreich
```

## Test 3 - Beachtung der Randfälle

<!-- grade_id: cell-71d2c10fcd3a87ca, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 10
building_type = "townhouse"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "A")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein auch wenn das Reihenhaus genau 10 Jahre alt ist.")
    print(e)
    
    
age = 15
building_type = "detached_house"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "B")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"B\" sein auch wenn das Einfamilienhaus genau 15 Jahre alt ist.")
    print(e)

age = 20
building_type = "detached_house"
insurance_class = insurance_classification(age, building_type)
try:
    assert_equal(insurance_class, "C")
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"C\" sein auch wenn das Einfamilienhaus genau 20 Jahre alt ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```