# Aufgabe 2 

Im Rahmen der operativen Planung in einem Handelsunternehmen soll erneut auf der Basis mehrerer gegebener
Faktoren eine Einschätzung bezüglich etwaiger Nachbestellungen für ein Produkt getroffen werden. 

Ist die Lagermenge eines bestimmten Produktes unter oder gleich dem Meldebestand, so soll dieses Produkt 200 mal nachbestellt werden.  
Ist die Lagermenge jedoch noch über dem Meldebestand so wollen Sie nichts nachbestellen. Geben Sie zusätzlich "keine Nachbestellung" mittels print() in die Konsole aus.

Implementieren Sie nun die Funktion `reorder`.  
Die Funktion erhält als Parameter `sales_yesterday` einen Integer Wert, welcher angibt wie oft ein Produkt am Vortag verkauft wurde.  
Sowie `stock_yesterday`, welcher angibt wie oft das Produkt gestern vor Verkaufsbeginn noch im Lager vorrätig war.  
Aus diesen beiden Werten können Sie sich den heutigen Lagerbestand berechnen.  
Der dritte Parameter ist `reorder_level`, also der Meldebestand. Ist dieser erreicht soll wie oben beschrieben nachbestellt werden.  
Als Ergebnis soll ein neuer Integer Wert `reorder_quantity` erstellt und zurückgegeben werden, welcher angibt wie viele Exemplare des Produktes nachbestellt werden sollen. Soll es nicht nachbestellt werden, so ist die `reorder_quantity` = 0.

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-ed147248e1bd5839, points: 0, grade: false, solution: true -->
```python
def reorder(reorder_level, stock_yesterday, sales_yesterday):
    
    reorder_quantity = 0
    stock_today = stock_yesterday - sales_yesterday
    if # Bedingung:
        reorder_quantity = 200
    else:
        print("keine Nachbestellung")
    
    return reorder_quantity


print("Wenn Ihre Funktion mit den Parametern \"reorder_level = 1000\", \"stock_yesterday = 1000\", \"sales_yesterday = 1500\" aufgerufen wird ist das Ergebnis:", reorder(1000, 1000, 1500))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "reorder_level = 1000", "stock_yesterday = 1000", "sales_yesterday = 1500" aufgerufen wird ist das Ergebnis: 200
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


```python
try:
    reorder(1000, 1000, 1500)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingung
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll anders sein je nachdem ob der Meldebestand erreicht ist oder nicht.

<!-- grade_id: cell-b5f89c14fd16f871, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_not_equal

reorder_level = 1000
reorder_level2 = 15000
sales_yesterday = 8000
stock_yesterday = 20000
reorder_quantity = reorder(reorder_level, stock_yesterday, sales_yesterday)
reorder_quantity2 = reorder(reorder_level2, stock_yesterday, sales_yesterday)
try:
    assert_not_equal(reorder_quantity, reorder_quantity2)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll unterschiedlich sein, je nachdem ob der Meldebestand erreicht wurde.")
    print(e)
    
```

Output:
```
keine Nachbestellung
Test 1 erfolgreich
```

## Test 2 - Korrekte Berechnung der Bestellmenge
Die Menge an zu bestellenden Produkten (`reorder_quantity`) soll entweder genau so groß sein wie die Verkaufsmenge am Vortag oder doppelt so groß.

<!-- grade_id: cell-8b34857b5ddbd553, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

reorder_level = 1000
sales_yesterday = 1000
stock_yesterday = 1500
reorder_quantity = reorder(reorder_level, stock_yesterday, sales_yesterday)
try:
    assert_equal(reorder_quantity, 200)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 200 sein.")
    print(e)
    
    
reorder_level = 1000
sales_yesterday = 200
stock_yesterday = 1500
reorder_quantity = reorder(reorder_level, stock_yesterday, sales_yesterday)
try:
    assert_equal(reorder_quantity, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein.")
    print(e)
    
```

Output:
```
Test 1 erfolgreich
keine Nachbestellung
Test 2 erfolgreich
```