# Aufgabe 4

Im Rahmen der operativen Planung in einem Handelsunternehmen soll erneut auf der Basis mehrerer gegebener
Faktoren eine Einschätzung bezüglich etwaiger Nachbestellungen für ein Produkt getroffen werden. 

Ist die Lagermenge eines bestimmten Produktes unter oder gleich dem Meldebestand, so muss dieses Produkt nachbestellt werden (1000 Stück).  
Sollte ein Produkt durch manuellen Wunsch nachbestellt werden müssen, so werden wie gewohnt nur 500 Stück bestellt.  
In beiden Fällen muss nun jedoch auch ausreichend Guthaben auf dem Geschäftskonto vorhanden sein. Das Produkt kostet pro Stück 2.25€. Sollte das Guthaben nicht für 1000 Stück reichen, so soll natürlich trotzdem geprüft werden ob eine manuelle Bestellung vorliegt und das Guthaben für diese reicht. *Nutzen Sie hierzu komplexe Bedingungen. Sie brauchen keine Verschachtelung*

Implementieren Sie nun die Funktion `bestellmenge_berechnen`. Die Funktion erhält als Parameter `meldebestand` einen Integer Wert, welcher den Meldebestand angibt.
Sowie `bestand_heute`, welcher angibt wie oft das Produkt an diesem Morgen noch im Lager vorrätig ist.
Der dritte Parameter `bestellwunsch` ist vom Typ Boolean und ist sobald ein manuel eingetragener Nachbestellwunsch vorliegt = True.  
Hinzu kommt ein Float Wert `guthaben`. Dieser speichert das Guthaben auf dem Geschäftskonto.  
Als Ergebnis soll eine neuer Integer Wert `bestellmenge` erstellt und zurückgegeben werden, welcher angibt wie viele Exemplare des Produktes nachbestellt werden sollen. Soll es nicht nachbestellt werden, so ist die `bestellmenge` = 0.

## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandeln Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-9e6b5dd1ea91a73e, points: 0, grade: false, solution: true -->
```python
def bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben):
    
    ###START
    ###BEGIN SOLUTION
    bestellmenge = 0
    if bestand_heute <= meldebestand and guthaben >= 1000 *2.25:
        bestellmenge = 1000
    elif bestellwunsch and guthaben >= 500 * 2.25:
        bestellmenge = 500
    
    ###END SOLUTION
    ###ENDE
    return bestellmenge

print("Wenn Ihre Funktion mit den Parametern \"meldebestand = 500\", \"bestand_heute = 499\", \"bestellwunsch = False\", \"guthaben = 99999\" aufgerufen wird ist das Ergebnis:", bestellmenge_berechnen(500, 499, False, 99999))
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-97b795e035f960c8, points: 0, grade: true, solution: false -->
```python
try:
    bestellmenge_berechnen(500, 499, False, 99999)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingungen bei genug Guthaben
Die Menge an zu bestellenden Produkten (`bestellmenge`) soll 1000 sein wenn der Meldebestand erreicht ist.  
Die Menge soll 500 sein wenn ein manueller Bestellwunsch vorliegt.  
Vorausgesetzt es ist für beides genug Guthaben verfügbar.

<!-- grade_id: cell-381ac45105f05f51, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

guthaben = 999999
meldebestand = 500
bestand_heute = 499
bestellwunsch = False
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 1000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde")
    print(e)
    
    
meldebestand = 500
bestand_heute = 501
bestellwunsch = False
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn der Meldebestand nicht erreicht wurde und keine manueller Bestellwunsch geäußert wurde")
    print(e)
    
    
meldebestand = 500
bestand_heute = 501
bestellwunsch = True
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 0)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 500 sein wenn ein manueller Bestellwunsch vorliegt aber der Meldebestand nicht erreicht wurde.")
    print(e)
    

meldebestand = 500
bestand_heute = 499
bestellwunsch = True
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 1000)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde, unabhängig davon ob ein Bestellwunsch vorliegt")
    print(e)
```

## Test 2 - Korrekte Bedingungen mit begrenztem Guthaben
Die Menge an zu bestellenden Produkten (`bestellmenge`) soll falls das Guthaben nicht reicht 0 sein.  
Sollte das Guthaben nur für die große Bestellung nicht reichen aber eine manuelle Bestellung vorliegen, so sollen 500 Stück bestellt werden.

<!-- grade_id: cell-6c6e3aadbcf7b04d, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

guthaben = 3000
meldebestand = 500
bestand_heute = 499
bestellwunsch = False
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 1000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 1000 sein wenn der Meldebestand erreicht wurde und genug Guthaben vorhanden ist")
    print(e)
    
    
guthaben = 500
meldebestand = 500
bestand_heute = 499
bestellwunsch = False
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 0)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 0 sein wenn der Meldebestand erreicht wurde aber zu wenig Guthaben vorhanden ist")
    print(e)
    

guthaben = 1500
meldebestand = 500
bestand_heute = 499
bestellwunsch = True
bestellmenge = bestellmenge_berechnen(meldebestand, bestand_heute, bestellwunsch, guthaben)
try:
    assert_equal(bestellmenge, 500)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 500 sein wenn das Guthaben zwar nicht für 1000Stück reicht aber für 500 und eine manuelle Bestellung vorliegt")
```