 # Aufgabe 3

 In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. In diesem Schritt wollen Sie die beiden vorherigen kombinieren. Ein Gebäude soll wie bereits in Aufgabe 1 immer dann in Risikoklasse "D" eingestuft werden wenn es sich um ein Mehrfamilienhaus handelt. Handelt es sich bei dem Gebäude jedoch nicht um ein Mehrfamilienhaus, so soll das Alter überprüft werden. Ist das Alter eines Gebäudes unter oder gleich 10 Jahren, so gehört es in die Risikoklasse "A".
 Ist ein Gebäude jedoch kein Mehrfamilienhaus aber trotzdem älter als 10 Jahre so ist die Risikoklasse "C".

 Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren). Sowie den Parameter `is_apartement_building`, der in Form eines Boolean-Wertes angibt ob das Gebäude ein Mehrfamilienhaus ist (True oder False).
 Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden.


 ### Hilfe
 Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
 Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-ede7e86ba6a80dd1, points: 0, grade: false, solution: true -->
```python
def insurance_classification(age, is_apartement_building):
    
    insurance_class = "A"
    if #Bedingung (wenn das Gebäude ein Mehrfamilienhaus ist):
        insurance_class = "D"
    elif #andere Bedingung (wenn das Gebäude unter oder gleich 10 Jahren alt ist):
        insurance_class = "A"
    else:
        insurance_class = "C"
        
    return insurance_class


print("Wenn Ihre Funktion mit den Parametern \"age = 5\" und \"is_apartement_building = True\" aufgerufen wird ist das Ergebnis:", insurance_classification(5, True))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "age = 5" und "is_apartement_building = True" aufgerufen wird ist das Ergebnis: D
```

 ***
 # Automatische Tests
 Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-c38b21e8b04ff752, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(5, True)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

 ## Test 1 - Korrekte Bedingungen
 Die Risikoklasse (`insurance_class`) soll "D" sein, wenn es sich bei dem Gebäude um ein Mehrfamilienhaus handelt.
 Ist dies nicht der Fall so soll die Risikoklasse "C" oder "A" sein, je nachdem ob das Gebäude älter als 10 Jahre ist (C) oder nicht (A).

<!-- grade_id: cell-6174197895da673b, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 5
is_apartement_building = True
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
    print(e)
    
    
age = 20
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "C")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"C\" sein wenn das Gebäude kein Mehrfamilienhaus ist aber schon älter als 10 Jahre.")
    print(e)
    

age = 5
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude kein Mehrfamilienhaus ist und noch nicht älter als 10 Jahre.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

 ## Test 2 - Beachtung der Randfälle
 Die Risikoklasse (`insurance_class`) soll auch dann "A" sein wenn das Gebäude kein Mehrfamilienhaus ist und *genau* 10 Jahre alt ist. Jedoch ab 11 Jahren soll die Klasse "C" sein.

<!-- grade_id: cell-5137ab65d8276eab, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 10
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein auch wenn das Gebäude genau genau 10 Jahre alt ist.")
    print(e)    
    
    
age = 11
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "C")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"C\" sein auch wenn das Gebäude nur knapp über 10 Jahre alt ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```

   ## Test 3 - Truthyness in Python
   In Python wird unter bestimmten Umständen auch andere Datentypen als `Boolean` als `True` oder `False`interpretiert.
   Dieser Test prüft ob auch wirklich nur die Risikoklasse D ausgegeben wird wenn ìs_apartment_building = True` gilt.

<!-- grade_id: cell-b741a1d5d93943be, points: 0, grade: true, solution: false -->
```python
age = 5
is_apartement_building = "Tolles Haus mit Aussicht"
insurance_class = insurance_classification(age, is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
```