  # Aufgabe 1

  In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Im ersten Schritt wollen Sie zunächst nur überprüfen ob es sich bei einem Gebäude um ein Mehrfamilienhaus handelt. Ist dies der Fall so ist die Risikoklasse in jedem Fall "D".

  Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `is_apartement_building`, welcher in Form eines Boolean-Wertes angibt ob das Gebäude ein Mehrfamilienhaus ist (True oder False).
  Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden. Ist das Gebäude ein Mehrfamilienhaus, so geben Sie "D" zurück. Ansonsten geben Sie "A" zurück.


## Wichtige Hinweise
Bisher haben Sie im Rahmen dieser Einführung in Python Funktionen noch nicht kennen gelernt. Aus diesem Grund hier ein paar Hinweise zum Umgang mit Funktionen im Rahmen dieser Übungsaufgaben:

1. Erstellen Sie Ihre Lösung innerhalb des mit `### START` und `### ENDE` markierten Bereichs. Ändern Sie den Programmcode außerhalb des markierten Bereichs nicht. 
1. Funktionen besitzen Parameter. Diese Parameter haben Namen. Behandel Sie die Parameter einfach wie Variablen in Ihrer Lösung. 
1. Bitte löschen Sie ggf die Zeile `raise NotImplementedError()`

### Hilfe
  Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
  Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-6ab468750feb45ab, points: 0, grade: true, solution: true -->
```python
def insurance_classification(is_apartement_building):
    
    ###START
    insurance_class = "A"
    if #Bedingung (wenn das Gebäude ein Mehrfamilienhaus ist):
        insurance_class = "D"
    
    ###ENDE
    return insurance_class


print("Wenn Ihre Funktion mit einem Parameter \"is_apartement_building = True\" aufgerufen wird ist das Ergebnis:", insurance_classification(True))
```

Output:
```
Wenn Ihre Funktion mit einem Parameter "is_apartement_building = True" aufgerufen wird ist das Ergebnis: D
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-d7d9f738cba20cc4, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(True)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

  ## Test 1 - Korrekte Bedingung
  Die Risikoklasse (`insurance_class`) soll "D" sein wenn das Gebäude ein Mehrfamilienhaus ist. Ansonsten soll sie "A" sein.

<!-- grade_id: cell-49278b35ff89383a, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

is_apartement_building = True
insurance_class = insurance_classification(is_apartement_building)
try:
    assert_equal(insurance_class, "D")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
    print(e)
    
    
is_apartement_building = False
insurance_class = insurance_classification(is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude kein Mehrfamilienhaus ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```

  ## Test 2 - Truthyness in Python
  In Python wird unter bestimmten Umständen auch andere Datentypen als `Boolean` als `True` oder `False`interpretiert.
  Dieser Test prüft ob auch wirklich nur die Risikoklasse D ausgegeben wird wenn ìs_apartment_building = True` gilt.

<!-- grade_id: cell-1b45c0887df397fa, points: 0, grade: true, solution: false -->
```python
is_apartement_building = "Tolles Haus mit Aussicht"
insurance_class = insurance_classification(is_apartement_building)
try:
    assert_equal(insurance_class, "A")
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll \"A\" sein wenn das Gebäude kein Mehrfamilienhaus ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
```