# Aufgabe 3 

In einer Bank sollen Kreditkartenanträge zukünftig mittels eines Computerprogramms automatisch bearbeitet werden. Sie sollen nun das Regelwerk zur Entscheidung über Anträge mittels `Conditionals` in einem Python Programm umsetzen.  
Im diesem Schritt wollen Sie die beiden vorherigen Schritte kombinieren und zusätzlich ein Limit für Kreditkarten festlegen, abhängig vom monatlichen Einkommen des Antragstellers.  
Prüfen Sie zunächst ob der Antragsteller unter 18 Jahren alt ist. Ist dies der Fall so soll keine Kreditkarte ausgestellt werden. Andernfalls soll geprüft werden ob der Antragsteller mehr als 3000€ verdient. In diesem Fall soll eine Kreditkarte ausgestellt werden mit einem Limit von 4000€. Ist dies nicht der Fall, so soll trotzdem geprüft werden ob der Antragsteller mehr als 1500€ verdient, damit eine Kreditkarte mit einem Limit von 2000€ ausgestellt werden kann. Andernfalls wird keine Kreditkarte ausgestellt.

Implementieren Sie nun die Funktion `request_verification`. Diese Funktion erhält einen Parameter `monthly_income`, welcher in Form eines float-Wertes das monatliche Einkommen des Antragstellers speichert. Sowie einen zweiten Parameter `age`, welcher in Form eines integer-Wertes das Alter des Antragstellers speichert.  
Als Ergebnis soll in Form eines Integer-Wertes das Kartenlimit zurückgegeben werden. Ist der Antragsteller unter 18 Jahren alt so setzen Sie das Limit auf -1 um deutlich zu machen, dass keine Karte ausgestellt werden soll.  
**Geben Sie zusätzlich mittels print() eine Meldung in die Konsole aus ob der Antrag abgelehnt oder angenommen wurde.**



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-11685dc7b37211f7, points: 0, grade: false, solution: true -->
```python
def request_verification(monthly_income, age):
    
    limit = 0
    if #Prüfen ob Person jünger als 18 Jahre ist:
        limit = -1
        print("Antrag abgelehnt")
    elif #Prüfen ob EInkommen über 3000€ ist:
        limit = 4000
        print("Antrag angenommen")
    elif #Prüfen ob Einkommen über 1500€ ist:
        limit = 2000
        print("Antrag angenommen")
    else:
        limit = -1
        print("Antrag abgelehnt")
        
    return limit


print("Wenn Ihre Funktion mit den Parametern \"monthly_income = 500\" und \"age = 25\" aufgerufen wird ist das Ergebnis:", request_verification(500, 25))
```

Output:
```
Antrag abgelehnt
Wenn Ihre Funktion mit den Parametern "monthly_income = 500" und "age = 25" aufgerufen wird ist das Ergebnis: -1
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-746d2425f28cb0c6, points: 0, grade: true, solution: false -->
```python
try:
    request_verification(500, 25)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

Output:
```
Antrag abgelehnt
```

## Test 1 - Korrekte Bedingungen
Der Rückgabewert (`limit`) soll *2000* betragen, falls der Antragsteller über 18 Jahren alt ist und über 1500€ monatlich verdient. Sollte der Antragsteller über 3000€ verdienen und über 18 Jahren alt sein, so soll der Rückgabewert *4000* betragen. Andernfalls soll der Rückgabewert *-1* sein.

<!-- grade_id: cell-f074918e652fa2b5, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 2500
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 2000)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 2000 betragen wenn der Antragsteller ein Einkommen von über 1500€ hat und über 18 Jahren alt ist.")
    print(e)
    
    
monthly_income = 3500
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 4000)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 4000 betragen wenn der Antragsteller ein Einkommen von über 3000€ hat und über 18 Jahren alt ist.")
    print(e)
    
    
monthly_income = 3500
age = 17
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller unter 18 Jahren alt ist. Dies ist unabhängig von seinem Einkommen.")
    print(e)
    
    
monthly_income = 1600
age = 17
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller unter 18 Jahren alt ist. Dies ist unabhängig von seinem Einkommen.")
    print(e)
    
    
monthly_income = 1300
age = 20
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 5 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller weniger als 1500€ monatlich verdient. Dies ist unabhängig von seinem Alter.")
    print(e)
    
    
```

Output:
```
Antrag angenommen
Test 1 erfolgreich
Antrag angenommen
Test 2 erfolgreich
Antrag abgelehnt
Test 3 erfolgreich
Antrag abgelehnt
Test 4 erfolgreich
Antrag abgelehnt
Test 5 erfolgreich
```

## Test 2 - Beachtung von Randfällen
Der Rückgabewert (`limit`) soll auch dann *-1* sein wenn der Antragsteller ein monatliches Einkommen von genau 1500€ hat. Jedoch soll der Wert schon *2000* betragen wenn das Einkommen nur minimal über 1500€ liegt.

<!-- grade_id: cell-ece4c74b4f1c6278, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 1500
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann -1 sein wenn der Antragsteller ein Einkommen von genau 1500€ hat.")
    print(e)
    
    
monthly_income = 1500.01
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 2000)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann 2000 betragen wenn der Antragsteller ein Einkommen von minimal mehr als 1500€ hat und über 18 Jahren alt ist.")
    print(e)


```

Output:
```
Antrag abgelehnt
Test 1 erfolgreich
Antrag angenommen
Test 2 erfolgreich
```