# Aufgabe 4 

In einer Bank sollen Kreditkartenanträge zukünftig mittels eines Computerprogramms automatisch bearbeitet werden. Sie sollen nun das Regelwerk zur Entscheidung über Anträge mittels `Conditionals` in einem Python Programm umsetzen.  
Im diesem Schritt wollen Sie folgende Regeln implementieren:  
- Falls der Antragsteller jünger als 18 Jahre alt ist wird keine Kreditkarte ausgestellt. 
- Falls der Antragsteller zwischen 18 und 30 Jahre alt ist wird nur eine Kreditkarte ausgestellt, wenn das Einkommen über 1500€ pro Monat liegt 
  - Ist der Antragsteller zwischen 18 und 30 und sein Einkommen liegt über 1500€ pro Monat wird eine Kreditkarte mit einem Limit von 2000€ ausgestellt. 
- Ist der Antragsteller zwischen 31 und 65 und sein Einkommen liegt über 2000€ pro Monat wird eine Kreditkarte mit einem Limit von 3000€ ausgestellt.  
- In allen anderen Fällen werden keine Kreditkarten ausgestellt

Implementieren Sie nun die Funktion `request_verification`. Diese Funktion erhält einen Parameter `monthly_income`, welcher in Form eines float-Wertes das monatliche Einkommen des Antragstellers speichert. Sowie einen zweiten Parameter `age`, welcher in Form eines integer-Wertes das Alter des Antragstellers speichert.  
Als Ergebnis soll in Form eines Integer-Wertes das Kartenlimit zurückgegeben werden. Falls keine Kreditkarte ausgestellt werden soll so setzen Sie das Limit auf -1 um dies deutlich zu machen.



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-1b3e7cef6d27ea43, points: 0, grade: false, solution: true -->
```python
def request_verification(monthly_income, age):
    
    ###BEGIN SOLUTION
    limit = 0
    if age < 18:
        limit = -1
    elif age < 31 and monthly_income > 1500:
        limit = 2000
    elif age <= 65 and monthly_income > 2000:
        limit = 3000
    else:
        limit = -1
        
    ###END SOLUTION
    return limit

print("Wenn Ihre Funktion mit den Parametern \"monthly_income = 500\" und \"age = 25\" aufgerufen wird ist das Ergebnis:", request_verification(500, 25))
```

Output:
```
Wenn Ihre Funktion mit den Parametern "monthly_income = 500" und "age = 25" aufgerufen wird ist das Ergebnis: -1
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-52b5489f22d7f7cd, points: 0, grade: true, solution: false -->
```python
try:
    request_verification(500, 25)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingungen

<!-- grade_id: cell-bd71deacca117a3d, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 2000
age = 17
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller unter 18 Jahren alt ist. Dies ist unabhägig vom Einkommen")
    print(e)

    
monthly_income = 1200
age = 40
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller ein Einkommen von weniger als 1500€ hat. Dies ist unabhägig vom Alter")
    print(e)
    
    
monthly_income = 4000
age = 20
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 2000)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 2000 betragen wenn der Antragsteller zwischen 18 und 30 Jahren alt ist und mehr als 1500€ verdient.")
    print(e)
    
    
monthly_income = 4000
age = 42
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 3000)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll 3000 betragen wenn der Antragsteller zwischen 31 und 65 Jahren alt ist und über 2000€ verdient.")
    print(e)
    
    
monthly_income = 4000
age = 70
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 5 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll -1 sein wenn der Antragsteller über 65 Jahren alt ist. Dies ist unabhängig vom Einkommen")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
Test 5 erfolgreich
```

## Test 2 - Beachtung von Randfällen
Der Rückgabewert (`limit`) soll auch dann *-1* sein wenn der Antragsteller ein monatliches Einkommen von genau 1500€ hat. Jedoch soll der Wert schon *2000* betragen wenn das Einkommen nur minimal über 1500€ liegt.

<!-- grade_id: cell-5e7c715a25064438, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

monthly_income = 1500
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, -1)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann -1 sein wenn der Antragsteller ein Einkommen von genau 1500€ hat.")
    print(e)
    
    
monthly_income = 1500.01
age = 25
limit = request_verification(monthly_income, age)
try:
    assert_equal(limit, 2000)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll auch dann 2000 betragen wenn der Antragsteller ein Einkommen von minimal mehr als 1500€ hat und über 18 Jahren alt ist.")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```