 # Aufgabe 1

 In einer Bank sollen Kreditkartenanträge zukünftig mittels eines Computerprogramms automatisch bearbeitet werden. Sie sollen nun das Regelwerk zur Entscheidung über Anträge mittels `Conditionals` in einem Python Programm umsetzen.
 In diesem ersten Schritt soll überprüft werden ob ein Antragsteller jünger als 18 Jahren alt ist. In diesem Fall soll keine Kreditkarte ausgestellt werden. Alle Antragsteller ab 18 Jahren sind berechtigt eine Kreditkarte zu erhalten.

 Implementieren Sie nun die Funktion `request_verification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines integer-Wertes das Alter des Antragstellers speichert.
 Als Ergebnis soll in Form eines Boolean-Wertes zurückgegeben werden ob der Antragsteller berechtigt ist eine Kreditkarte zu erhalten. Ist der Antragsteller älter als 18 Jahre so geben Sie *True* zurück. Ansonsten *False*.



 ### Hilfe
 Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
 Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-733b43fac6e9ea69, points: 0, grade: false, solution: true -->
```python
def request_verification(age):
    
    is_eligible = True
    if #Bedingung:
        #umstellen des Rückgabewertes
        
    return is_eligible


print("Wenn Ihre Funktion mit dem Parameter \"age = 25\" aufgerufen wird ist das Ergebnis:", request_verification(25))
```

Output:
```
Wenn Ihre Funktion mit dem Parameter "age = 25" aufgerufen wird ist das Ergebnis: True
```

 ***
 # Automatische Tests
 Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-14c7df1478e0fc44, points: 0, grade: true, solution: false -->
```python
try:
    request_verification(25)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

 ## Test 1 - Rückgabe als Boolean-Wert
 Der Rückgabewert (`is_eligible`) soll vom Datentypen Boolean sein.

<!-- grade_id: cell-d7a048039b2238aa, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 15
is_eligible = request_verification(age)
try:
    assert_equal(isinstance(is_eligible, bool), True)
    assert_equal(isinstance(is_eligible, str), False)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll vom Typen Boolean sein.\nDas Ergebnis soll nicht vom Typen String sein. Nutzen Sie keine Anführungszeichen für die Zuweisung von Boolean-Werten")
    print(e)
    
    
```

Output:
```
Test 1 erfolgreich
```

 ## Test 2 - Korrekte Bedingung
 Der Rückgabewert (`is_eligible`) soll *False* sein wenn der Antragsteller unter 18 Jahren alt ist. Ansonsten soll der Wert *True* sein.

<!-- grade_id: cell-b2e0e70e9c95453f, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 15
is_eligible = request_verification(age)
try
    assert_equal(is_eligible, False)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll False sein wenn der Antragsteller unter 18 Jahren alt ist.")
    print(e)
    
    
age = 20
is_eligible = request_verification(age)
try:
    assert_equal(is_eligible, True)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll True sein wenn der Antragsteller über 18 Jahren alt ist.")
    print(e)
    
    
```

 ## Test 3 - Beachtung von Randfällen
 Der Rückgabewert (`is_eligible`) soll auch dann *True* sein wenn der Antragsteller genau 18 Jahre alt ist.

<!-- grade_id: cell-ae00b6e0da823905, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 18
is_eligible = request_verification(age)
try:
    assert_equal(is_eligible, True)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis soll True sein wenn der Antragsteller genau 18 Jahre alt ist.")
    print(e)
    
```

Output:
```
Test 1 erfolgreich
```