# Aufgabe 4

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Analyse der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden. So soll aus einer Liste mit vielen Einnahmen heruasgefiltert werden, welches die höchste eingenommene Umsatzsteuer ist. Außerdem soll die Summe aller eingenommenen Umsatzsteuern berechnet werden.

Implementieren Sie nun die einzelnen Funktionen `calculate_taxes` und `calculate_highest`.  
Diese Funktionen erhalten als Parameter jeweils eine Liste `sales`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).   
Als Ergebnis sollen (je nach Funktion):  
Der gesamte Betrag, welcher durch die Umsatzsteuer eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden ("taxes").  
Der höchste Betrag, welcher durch die Umsatzsteuer an einem Artikel eingenommen wurde, in Form eines Float-Wertes ("highest_tax").    

Des Weiteren soll eine Hauptfunktion (`analyse_taxes`) implementiert werden. Diese soll die oben genannten Funktionen nur aufrufen und das Ergebnis ausgeben. Dazu soll auch die Hauptfunktion als Parameter eine Liste `sales` akzeptieren. Die Ausgabe soll wie folgt aussehen:  
`Die insgesamt eingenommene Umsatzsteuer beträgt: 180.7413
Die höchste eingenommene Umsatzsteuer beträgt: 95.0`


Achten Sie bei dieser Aufgabe darauf dass sowohl alle Funktionen einzeln funktionieren, als auch die gesamte Ausgabe über die Hauptfunktion (analyse_taxes).

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $


### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-d430232ec3d29238, points: 0, grade: false, solution: true -->
```python
# Hauptfunktion

def calculate_taxes(sales):
    taxes = 0
    for sale in sales:
        taxes += sale * 0.19
    return taxes
    
def calculate_highest(sales):
    highest_tax = 0
    for sale in sales:
        if (sale * 0.19) > highest_tax:
            highest_tax = sale * 0.19
    return highest_tax
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-a1204d9bf4dfef11, points: 0, grade: true, solution: false -->
```python
try:
    analyse_taxes([299.99, 250.00, 23.00])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

Output:
```
Die insgesamt eingenommene Umsatzsteuer beträgt: 108.8681
Die höchste eingenommene Umsatzsteuer beträgt: 56.9981
```

## Test 1 - Korrekte Berechnung der gesamten Umsatzsteuer
Die Umsatzsteuer (`taxes`) soll mittels der Funktion `calculate_taxes` korrekt berechnet werden.

<!-- grade_id: cell-940f7131dcf9bda6, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [519.55, 250.00, 19.37, 12.34, 234.56]
taxes = calculate_taxes(sales)
assert_equal(taxes, 196.80579999999998, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

sales = [99999999.99, 1234567.89, 9876543.22]
taxes = calculate_taxes(sales)
assert_equal(taxes, 21111111.108999997, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

sales = [12.34, 20.00]
taxes = calculate_taxes(sales)
assert_equal(taxes, 6.1446, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Bestimmung der höchsten Umsatzsteuer
Die höchste Umsatzsteuer (`highest_tax`) soll mittels der Funktion `calculate_highest` für jede Liste korrekt bestimmt werden.  

<!-- grade_id: cell-c6675d6074a2b535, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300.00, 250.00, 23.00, 210.00, 234.00]
balance = calculate_highest(sales)
assert_equal(balance, 57, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

sales = [200.00, 200.00, 200.00, 200.00]
balance = calculate_highest(sales)
assert_equal(balance, 38, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

sales = [15.00]
balance = calculate_highest(sales)
assert_equal(balance, 2.85, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

sales = [300.00, 500.00, 23.00, 210.00, 234.00]
balance = calculate_highest(sales)
assert_equal(balance, 95, "Das Ergebnis ist falsch")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```

## Test 3 - Korrekte Ausgabe
Die Hauptfunktion `analyse_taxes` soll die beiden oben getesteten Funktionen aufrufen und die Ausgabe erstellen wie im Aufgabentext beschrieben.  
**Achtung: Der Ausgabetext muss exakt mit dem aus der Aufgabenstellung übereinstimmen damit dieser Test funktioniert**

<!-- grade_id: cell-dd66d9b6a2727a0d, points: 0, grade: true, solution: false -->
```python
from unittest import mock
from nose.tools import assert_equal
import unittest
import io

with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    analyse_taxes([299.99, 250.00, 23.00, 210.00, 234.00])

assert_equal(fake_stdout.getvalue(), 'Die insgesamt eingenommene Umsatzsteuer beträgt: 193.2281\nDie höchste eingenommene Umsatzsteuer beträgt: 56.9981\n', "Die Ausgabe ist nicht perfekt.")
print("Test 1 erfolgreich.")
```

Output:
```
Test 1 erfolgreich.
```