<div class="alert alert-block alert-warning">

# Aufgabe 1

In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Im zweiten Schritt wollen Sie überprüfen ob das Alter eines Gebäudes unter oder gleich 10 Jahren ist. Ist dies der Fall so ist die Risikoklasse in jedem Fall "A".  

Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren).  
Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden. Ist das Gebäude älter als 10 Jahre, so geben Sie "D" zurück. Ansonsten geben Sie "A" zurück. Zusätzlich geben Sie bitte mittels print() in die Konsole aus ob das Gebäude als "neuwertig" oder "alt" eingestuft wurde.
</div>
<div class="alert alert-block alert-info">

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).
</div>

<!-- grade_id: cell-b4793346c766032e, points: 0, grade: false, solution: true -->
```python

    
    insurance_class = "D"
    if age <= 10:
        insurance_class = "A"
        print("Eingestuft als neuwertig.")
    #Else-Fall
    
    
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-e174a2b35a81205d, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(20)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingung
Die Risikoklasse (`insurance_class`) soll "D" sein wenn das Gebäude älter als 10 Jahre ist. Ansonsten soll sie "A" sein.

<!-- grade_id: cell-67a7104f4b6c7489, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 20
insurance_class = insurance_classification(age)

assert_equal(insurance_class, "D", "Das Ergebnis soll \"D\" sein wenn das Gebäude älter als 10 Jahre ist.")
print("Test 1 erfolgreich")

age = 5
insurance_class = insurance_classification(age)

assert_equal(insurance_class, "A", "Das Ergebnis soll \"A\" sein wenn das Gebäude noch keine 10 Jahre alst ist.")
print("Test 2 erfolgreich")

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-2-ef3e768c74f5> in <module>
----> 1 from nose.tools import assert_equal
      2 
      3 age = 20
      4 insurance_class = insurance_classification(age)
      5 

ModuleNotFoundError: No module named 'nose'
```

## Test 2 - Beachtung der Randfälle
Die Risikoklasse (`insurance_class`) soll auch dann "A" sein wenn das Gebäude *genau* 10 Jahre alt ist. Jedoch ab 11 Jahren soll die Klasse "D" sein.

<!-- grade_id: cell-95aa7cd60a07f3c3, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 11
insurance_class = insurance_classification(age)

assert_equal(insurance_class, "D", "Das Ergebnis soll \"D\" sein auch wenn das Gebäude nur knapp älter als 10 Jahre ist.")
print("Test 1 erfolgreich")

age = 10
insurance_class = insurance_classification(age)

assert_equal(insurance_class, "A", "Das Ergebnis soll \"A\" sein auch wenn das Gebäude genau 10 Jahre alst ist.")
print("Test 2 erfolgreich")
```

Output:
```
Eingestuft als alt.
Test 1 erfolgreich
Eingestuft als neuwertig.
Test 2 erfolgreich
```