<div class="alert alert-block alert-warning">

# Aufgabe 4

Es soll ein Programm zur automatischen Analyse von Kontobewegungen implementiert werden. So soll aus einer großen Liste von Zugängen und Abgängen erkannt werden wie der neue Kontostand ist. Außerdem soll der größte Zu- oder Abgang bestimmt werden.

Implementieren Sie nun die einzelnen Funktionen `calculate_balance` und `calculate_highest`.  
Diese Funktionen erhalten als Parameter jeweils eine Liste `transfers`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).  
Als Ergebnis sollen (je nach Funktion):  
Der `Kontostand` in Form eines Float-Wertes zurückgegeben werden ("balance").  
Der `höchste Transfer` in Form eines Float-Wertes zurückgegeben werden ("highest_transfer").  

Des Weiteren soll eine Hauptfunktion (`analyse_transfers`) implementiert werden. Diese soll die oben genannten Funktionen nur aufrufen und das Ergebnis ausgeben. Dazu soll auch die Hauptfunktion als Parameter eine Liste `transfers` akzeptieren. Die Ausgabe soll wie folgt aussehen:  
`Der Kontostand beträgt: 300.05
Der höchste Transfer beträgt: 500.0`


Achten Sie bei dieser Aufgabe darauf dass sowohl alle Funktionen einzeln funktionieren, als auch die gesamte Ausgabe über die Hauptfunktion (analyse_transfers).
</div>
<div class="alert alert-block alert-info">

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).
</div>

<!-- grade_id: cell-1a07b590076476d0, points: 0, grade: false, solution: true -->
```python
#Hauptfunktion

#calculate_highest Funktion

def calculate_balance(transfers):
    balance = 0
    for transfer in transfers:
        balance += transfer
    return balance
    

```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-d6c1618d07297752, points: 0, grade: true, solution: false -->
```python
try:
    calculate_balance([+519.55, -250.00, -19.37])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung des Kontostands
Der Kontostand (`balance`) soll mittels der Funktion `calculate_balance` korrekt berechnet werden.

<!-- grade_id: cell-247d82df8690edfd, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+519.55, -250.00, -19.37, -12.34, +234.56]
balance = calculate_balance(transfers)
assert_equal(balance, 472.40, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+99999999.99, -1234567.89, -9876543.22]
balance = calculate_balance(transfers)
assert_equal(balance, 88888888.88, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [-12.34, +20.00]
balance = calculate_balance(transfers)
assert_equal(balance, 7.66, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Bestimmung der höchsten Bankbewegung
Die höchste Bankbewegung (`highest_transfer`) soll mittels der Funktion `calculate_highest` für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-69e846172462494b, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+300.00, -250.00, -23.00, +210.00, +234.00]
highest_transfer = calculate_highest(transfers)
assert_equal(highest_transfer, 300, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+200.00, -200.00, +200.00, -200.00]
highest_transfer = calculate_highest(transfers)
assert_equal(highest_transfer, 200, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [+15.00]
highest_transfer = calculate_highest(transfers)
assert_equal(highest_transfer, 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

transfers = [+300.00, -500.00, -23.00, +210.00, +234.00]
highest_transfer = calculate_highest(transfers)
assert_equal(highest_transfer, -500, "Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```

## Test 3 - Korrekte Ausgabe
Die Hauptfunktion `analyse_transfers` soll die beiden oben getesteten Funktionen aufrufen und die Ausgabe erstellen wie im Aufgabentext beschrieben.  
**Achtung: Der Ausgabetext muss exakt mit dem aus der Aufgabenstellung übereinstimmen damit dieser Test funktioniert**

<!-- grade_id: cell-40e2b6b56c4ead9b, points: 0, grade: true, solution: false -->
```python
from unittest import mock
from nose.tools import assert_equal
import unittest
import io

with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    analyse_transfers([+300.00, -250.00, -23.00, +210.00, +234.00])

assert_equal(fake_stdout.getvalue(), 'Der Kontostand beträgt: 471.0\nDer höchste Transfer beträgt: 300.0\n', "Die Ausgabe ist nicht perfekt.")
print("Test 1 erfolgreich.")
```

Output:
```
Test 1 erfolgreich.
```