# Aufgabe 3

Im Rahmen der Bestellungsabwicklung soll ein Programm zur automatischen Bewertung mehrerer Bestellungen implementiert werden. So soll anhand einer Liste von Bestellungen und einer (separaten) Liste von Zahlungen erkannt werden, welche Bestellungen noch nicht bezahlt wurden.  

Implementieren Sie nun die Funktion `analyse_orders`. Diese Funktion erhält als Parameter eine Liste `orders`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `orders` eine Liste bestehend aus der Bestellnummer als erstem Wert, gefolgt von einer Artikelbezeichnung. Die Liste sieht also in etwa so aus: `[[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], ["10003", "Früchtemüsli"]]`.   
Zusätzlich gibt es als zweiten Parameter eine weitere Liste `payments`. Diese Liste beinhält alle Bestellnummern, welche bezahlt wurden. Also in etwa so: `[10001, 10003]`.  
Als Ergebnis sollen alle Bestellungen zurückgegeben werden, welche noch nicht bezahlt wurden. Für die obigen Beispiele wäre die Ergebnis-Liste also: `[[10002, "Blaubeermüsli"]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-aa93af9f23515e63, points: 0, grade: false, solution: true -->
```python

    
    unpaid_orders = orders

    for #durchlaufe alle bestellungen: 
        for #durchlaufe alle bezahlungen:
            if payment == order[0]:
                unpaid_orders.remove(order)
                

```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-5af9ee7d975a6047, points: 0, grade: true, solution: false -->
```python
try:
    analyse_orders([[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"]], [10001, 10003])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Filterung der Bestellungen
Die noch *nicht* bezahlten Bestellungen sollen zurückgegeben werden.

<!-- grade_id: cell-2755d5398e49ab99, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"]]
payments = [10001, 10003]
paid_orders = analyse_orders(orders, payments)
assert_equal(paid_orders, [[10002, "Blaubeermüsli"]], "Die Bestellung \"10002 Blaubeermüsli\" soll zurückgegeben werden.")
print("Test 1 erfolgreich")

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"], [10004, "Früchtemüsli"]]
payments = [10001, 10003]
paid_orders = analyse_orders(orders, payments)
assert_equal(paid_orders, [[10002, "Blaubeermüsli"], [10004, "Früchtemüsli"]], "Die Bestellungen \"10002 Blaubeermüsli\" und \"10004 Früchtemüsli\" sollen zurückgegeben werden.")
print("Test 2 erfolgreich")

orders = [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"], [10004, "Früchtemüsli"]]
payments = [10004]
paid_orders = analyse_orders(orders, payments)
assert_equal(paid_orders, [[10001, "Erdbeermüsli"], [10002, "Blaubeermüsli"], [10003, "Früchtemüsli"]], "Die Bestellungen \"10001 Erdbeermüsli\", \"10002 Blaubeermüsli\" und \"10003 Früchtemüsli\" sollen zurückgegeben werden.")
print("Test 3 erfolgreich")

```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```