# Aufgabe 3

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Bewertung mehrerer Materialien implementiert werden. So sollen aus einer Liste mehrerer Materialien alle herausgefiltert werden, deren Lagerbestand unter dem Meldebstand von 200 liegt, damit diese neu bestellt werden können.  

Implementieren Sie nun die Funktion `analyse_stock`. Diese Funktion erhält als Parameter eine Liste `materials`. Diese Liste besteht allerding aus mehreren Listen. So ist jedes Element der Liste `materials` eine Liste bestehend aus dem Materialnamen als erstem Wert, gefolgt von dem Materialbestand als zweitem Wert. Die Liste sieht also in etwa so aus: `[["Schrauben", 330], ["Nägel", 2500], ["Muttern", 150]]`.  

Als Ergebnis soll eine Liste mit nur den Materialien unter dem Meldebstand von 200Stück zurückgegeben werden (`to_reorder`). Ergebnis für obiges Beispiel wäre also: `[["Muttern", 150]]`



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-35309a5761d24069, points: 0, grade: false, solution: true -->
```python
    
    
    to_reorder = []
    for material in materials:
        if material[1] < 200:
            # Ergänze die Liste der zu bestellenden Materialien durch das aktuell geprüfte
    
    
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-8d8c306f10e0bb1a, points: 0, grade: true, solution: false -->
```python
try:
    analyse_stock([["Schrauben", 330], ["Nägel", 2500]])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Richtige Filterung
Es sollen die Materialien herausgefiltert werden, bei denen der Bestand unter 200 Stück liegt.

<!-- grade_id: cell-bfba5cc6b7420fdf, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

materials = [["Schrauben", 330], ["Nägel", 2500], ["Muttern", 150]]
to_reorder = analyse_stock(materials)
assert_equal(to_reorder, [["Muttern", 150]], "Das Material \"Muttern\" sollte als einziges Material zurückgegeben werden")
print("Test 1 erfolgreich")

materials = [["Gewinde", 230], ["Knopfzellen", 120]]
to_reorder = analyse_stock(materials)
assert_equal(to_reorder, [["Knopfzellen", 120]], "Das Material \"Knopfzellen\" sollte als einziges Material zurückgegeben werden")
print("Test 2 erfolgreich")

materials = [["Mandeln", 150], ["Haselnüsse", 2500], ["Erdnüsse", 300], ["Haselnüsse", 180]]
to_reorder = analyse_stock(materials)
assert_equal(to_reorder, [["Mandeln", 150], ["Haselnüsse", 180]], "Die Materialien \"Mandeln\" und \"Haselnüsse\" sollten als zurückgegeben werden")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```