# Aufgabe 2

In einem Versicherungsunternehmen soll, im Rahmen der Gebäudeversicherungen, ein Programm zur automatischen Risikobewertung implementiert werden. Im vierten Schritt soll weiter zwischen den Risikoklassen differenziert werden. Ein Gebäude soll immer dann in Risikoklasse "D" eingestuft werden wenn es sich um ein Mehrfamilienhaus handelt *oder* es älter als 20 Jahre ist. Andernfalls soll folgendes überprüft werden:  
Ist das Alter eines Gebäudes unter oder gleich 10 Jahren, so gehört es in die Risikoklasse "A".  
Ist ein Gebäude im Alter von 11-20 Jahren, so ist die Risikoklasse "B". 

Implementieren Sie nun die Funktion `insurance_classification`. Diese Funktion erhält einen Parameter `age`, welcher in Form eines Integer-Wertes angibt wie alt das Gebäude ist (in Jahren). Sowie den Parameter `is_apartement_building`, der in Form eines Boolean-Wertes angibt ob das Gebäude ein Mehrfamilienhaus ist (True oder False).  
Als Ergebnis soll jedoch die Risikoklasse in Form eines Buchstabens zurückgegeben werden ("A","B" oder "D"). 


### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `Kontrollstrukturen`in Python finden Sie [hier](../../30_conditionals/10_conditionals.ipynb).

<!-- grade_id: cell-ca05203a84dbc45b, points: 0, grade: false, solution: true -->
```python

    
    insurance_class = "A"
    if #Bedingung wann Gebäude als Klasse "D" eingestuft werden soll:
        insurance_class = "D"
    elif # Bedingung für Klasse "A":
        insurance_class = "A"
    # Else-Fall
        
        
```

***
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-3f5d8b80f1189754, points: 0, grade: true, solution: false -->
```python
try:
    insurance_classification(21, True)
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bedingungen
Die Risikoklasse (`insurance_class`) soll "D" sein, wenn es sich bei dem Gebäude um ein Mehrfamilienhaus handelt *oder* das Gebäude älter als 20 Jahre ist.  
Ist dies nicht der Fall so soll die Risikoklasse "B" oder "A" sein, je nachdem ob das Gebäude jünger als 10 Jahre ist (A) oder zwischen 11 und 20 Jahren (B).

<!-- grade_id: cell-21d61db97fac3cca, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 5
is_apartement_building = True
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "D", "Das Ergebnis soll \"D\" sein wenn das Gebäude ein Mehrfamilienhaus ist.")
print("Test 1 erfolgreich")

age = 25
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "D", "Das Ergebnis soll \"D\" sein wenn das Gebäude älter als 20 Jahre ist.")
print("Test 2 erfolgreich")

age = 15
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "B", "Das Ergebnis soll \"B\" sein wenn das Gebäude kein Mehrfamilienhaus ist und zwischen 11 und 20 Jahren alt.")
print("Test 3 erfolgreich")

age = 5
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "A", "Das Ergebnis soll \"A\" sein wenn das Gebäude kein Mehrfamilienhaus ist und noch nicht älter als 10 Jahre.")
print("Test 4 erfolgreich")

age = 25
is_apartement_building = True
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "D", "Das Ergebnis soll \"D\" sein wenn das Gebäude sowohl ein Mehrfamilienhaus, als auch älter als 20 ist.")
print("Test 5 erfolgreich")

```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
Test 5 erfolgreich
```

## Test 2 - Beachtung der Randfälle
Die Risikoklasse (`insurance_class`) soll auch dann "A" sein wenn das Gebäude kein Mehrfamilienhaus ist und *genau* 10 Jahre alt ist. Ebenso gehört ein *genau* 20 Jahre altes Gebäude noch zur Risikoklasse "B".

<!-- grade_id: cell-544bf9ec81461f19, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

age = 10
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "A", "Das Ergebnis soll \"A\" sein auch wenn das Gebäude genau 10 Jahre alt ist.")
print("Test 1 erfolgreich")

age = 20
is_apartement_building = False
insurance_class = insurance_classification(age, is_apartement_building)

assert_equal(insurance_class, "B", "Das Ergebnis soll \"B\" sein auch wenn das Gebäude genau 20 Jahre alt ist.")
print("Test 2 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
```