<div class="alert alert-block alert-warning">

# Aufgabe 4

Es soll ein Programm zur automatischen Analyse von Materialeingängen und Materialabgängen implementiert werden. So soll aus einer großen Liste von Eingängen und Abgängen erkannt werden wie der neue Materialbestand ist. Außerdem soll der größte Ein- oder Abgang bestimmt werden.

Implementieren Sie nun die einzelnen Funktionen `calculate_stock` und `calculate_highest`.  
Diese Funktionen erhalten als Parameter jeweils eine Liste `incoming_outgoing`, welche in Form von Integer-Werten die Eingänge und Abgänge in das Materiallager angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis sollen (je nach Funktion):  
Der `Materialbestand` in Form eines Integer-Wertes zurückgegeben werden ("stock").  
Der `höchste Transfer` in Form eines Integer-Wertes zurückgegeben werden ("highest_transfer").  

Des Weiteren soll eine Hauptfunktion (`analyse_stock`) implementiert werden. Diese soll die oben genannten Funktionen nur aufrufen und das Ergebnis ausgeben. Dazu soll auch die Hauptfunktion als Parameter eine Liste `incoming_outgoing` akzeptieren. Die Ausgabe soll wie folgt aussehen:  
`Der Materialbestand beträgt: 305
Der höchste Transfer beträgt: 500`


Achten Sie bei dieser Aufgabe darauf dass sowohl alle Funktionen einzeln funktionieren, als auch die gesamte Ausgabe über die Hauptfunktion (analyse_stock).
</div>
<div class="alert alert-block alert-info">

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).
</div>

<!-- grade_id: cell-5e851adeef6de487, points: 0, grade: false, solution: true -->
```python
# Hauptfunktion

# calculate_highest Funktion

def calculate_stock(incoming_outgoing):
    stock = 0
    for transfer in incoming_outgoing:
        stock += transfer
    return stock
    
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-a459a1e684f8555b, points: 0, grade: true, solution: false -->
```python
try:
    analyse_stock([-19, -12, +234])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

Output:
```
Der Materialbestand beträgt: 203
Der höchste Transfer beträgt: 234
```

## Test 1 - Korrekte Berechnung des Materialbestands
Der Materialbestand (`stock`) soll mittels der Funktion `calculate_stock` korrekt berechnet werden.

<!-- grade_id: cell-73dfdf6644d28eb5, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 472, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 88888889, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [-12, +20]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 8, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Bestimmung der höchsten Materialbewegung
Die höchste Materialbewegung (`highest_transfer`) soll mittels der Funktion `calculate_highest` für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-9bc03493d6cfdd2c, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+300, -250, -23, +210, +234]
highest_transfer = calculate_highest(incoming_outgoing)
assert_equal(highest_transfer, 300, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+200, -200, +200, -200]
highest_transfer = calculate_highest(incoming_outgoing)
assert_equal(highest_transfer, 200, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [+15]
highest_transfer = calculate_highest(incoming_outgoing)
assert_equal(highest_transfer, 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

incoming_outgoing = [+300, -500, -23, +210, +234]
highest_transfer = calculate_highest(incoming_outgoing)
assert_equal(highest_transfer, -500, "Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```

## Test 3 - Korrekte Ausgabe
Die Hauptfunktion `analyse_stock` soll die beiden oben getesteten Funktionen aufrufen und die Ausgabe erstellen wie im Aufgabentext beschrieben.  
**Achtung: Der Ausgabetext muss exakt mit dem aus der Aufgabenstellung übereinstimmen damit dieser Test funktioniert**

<!-- grade_id: cell-aa498808ea48d927, points: 0, grade: true, solution: false -->
```python
from unittest import mock
from nose.tools import assert_equal
import unittest
import io

with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    analyse_stock([+300, -250, -23, +210, +234])

assert_equal(fake_stdout.getvalue(), 'Der Materialbestand beträgt: 471\nDer höchste Transfer beträgt: 300\n', "Die Ausgabe ist nicht perfekt.")
print("Test 1 erfolgreich.")
```

Output:
```
Test 1 erfolgreich.
```