<div class="alert alert-block alert-warning">

# Aufgabe 5

Es soll ein Programm zur automatischen Berechnung eines Kontostandes anhand von Zugängen und Abgängen auf ein Konto implementiert werden.  
**Allerdings sollen Sie hierbei auf Schleifen verzichten. Lösen Sie diese Aufgabe durch eine rekursive Funktion.  
Das heißt die Funktion soll sich immer wieder selbst aufrufen mit einem veränderten Parameter.**

Implementieren Sie nun die Funktion `calculate_balance`. Diese Funktion erhält als Parameter eine Liste `transfers`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).  
Als Ergebnis soll der Kontostand in Form eines Float-Wertes zurückgegeben werden.
</div>
<div class="alert alert-block alert-info">

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).
</div>

<!-- grade_id: cell-5ae53d3aec3098e8, points: 0, grade: false, solution: true -->
```python
### Dies ist lediglich ein Pseudocode. Sie müssen den Code komplett umschreiben, die Logik sollte jedoch erhalten bleiben

def calculate_balance(transfers):
    if (die Liste hat nur einen Eintrag):
        return (Eintrag aus der Liste)
    else:
        #
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-6c118496ffadc703, points: 0, grade: true, solution: false -->
```python
try:
    calculate_balance([+519.55, -250.00, -19.37])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Kontostand (`balance`) soll für jede Liste korrekt berechnet werden.  
**Diese Testfälle können nicht überprüfen ob Sie die Aufgabe rekursiv gelöst haben. Achten Sie selbst darauf dass Sie keine Schleifen in Ihrem Programm nutzen.**

<!-- grade_id: cell-45e429fa047a093a, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+519.55, -250.00, -19.37, -12.34, +234.56]
balance = calculate_balance(transfers)
assert_equal(balance, 472.40, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+99999999.99, -1234567.89, -9876543.22]
balance = calculate_balance(transfers)
assert_equal(balance, 88888888.88, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [-12.34, +20.00]
balance = calculate_balance(transfers)
assert_equal(balance, 7.66, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```