<div class="alert alert-block alert-warning">

# Aufgabe 5

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Berechnung der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden.  
**Allerdings sollen Sie hierbei auf Schleifen verzichten. Lösen Sie diese Aufgabe durch eine rekursive Funktion.  
Das heißt die Funktion soll sich immer wieder selbst aufrufen mit einem veränderten Parameter.**

Implementieren Sie nun die Funktion `calculate_taxes`. Diese Funktion erhält als Parameter eine Liste `sales`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).  
Als Ergebnis soll der gesamte Betrag, welcher durch die Umsatzsteuer eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden.   

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $
</div>
<div class="alert alert-block alert-info">

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).
</div>

<!-- grade_id: cell-050636cce2c9d626, points: 0, grade: false, solution: true -->
```python
### Dies ist lediglich ein Pseudocode. Sie müssen den Code komplett umschreiben, die Logik sollte jedoch erhalten bleiben

def calculate_taxes(sales):
    if (die Liste hat nur einen Eintrag):
        return (Eintrag aus der Liste)
    else:
        #
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-e9a955689d59bfc2, points: 0, grade: true, solution: false -->
```python
try:
    calculate_taxes([300, 250.45, 19.78])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Betrag (`taxes`) soll für jede Liste korrekt berechnet werden.  
**Diese Testfälle können nicht überprüfen ob Sie die Aufgabe rekursiv gelöst haben. Achten Sie selbst darauf dass Sie keine Schleifen in Ihrem Programm nutzen.**

<!-- grade_id: cell-753f028e9a0186c0, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

sales = [300, 250.45, 19.78, 123.56, 50.23]
taxes = calculate_taxes(sales)
assert_equal(taxes, 141.3638, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

sales = [99999999, 1234567, 9876543]
taxes = calculate_taxes(sales)
assert_equal(taxes, 21111110.709999997, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

sales = [12, 20]
taxes = calculate_taxes(sales)
assert_equal(taxes, 6.08, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```