# Aufgabe 2

Es soll ein Programm zur automatischen Analyse von Kontobewegungen implementiert werden. So soll aus einer großen Liste von Zugängen und Abgängen erkannt werden wie der neue Kontostand ist. Außerdem soll die durchschnittliche Höhe der Kontobewegungen berechnet werden, sowie der größte Zu- oder Abgang bestimmt werden.

Implementieren Sie nun die Funktion `analyse_transfers`. Diese Funktion erhält als Parameter eine Liste `transfers`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).  
Als Ergebnis soll der `Kontostand` als Float-Wert zurückgegeben werden ("balance").  
Die `durchschnittliche Höhe` einer Kontobewegung ("average").  
Der `höchste Transfer` ("highest_transfer"). Hierbei sollen sowohl negative als auch positive Transfers berücksichtigt werden.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-bdc7634046ec3fd1, points: 0, grade: false, solution: true -->
```python


    balance = 0
    highest_transfer = 0
    for transfer in transfers:
        # Rückgabewert "Balance" mit jedem durchlauf ändern
        if transfer ** 2 > highest_transfer ** 2:
            highest_transfer = transfer
    
    # Rückgabewert "average" ändern
    
    
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-3489e9a96afc91aa, points: 0, grade: true, solution: false -->
```python
try:
    analyse_transfers([+519.55, -250.00, -19.37])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung des Kontostands
Der Kontostand (`balance`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-bcd8176b0e96b318, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+519.55, -250.00, -19.37, -12.34, +234.56]
balance = analyse_transfers(transfers)
assert_equal(balance[0], 472.40, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+99999999.99, -1234567.89, -9876543.22]
balance = analyse_transfers(transfers)
assert_equal(balance[0], 88888888.88, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [-12.34, +20.00]
balance = analyse_transfers(transfers)
assert_equal(balance[0], 7.66, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Berechnung des Durchschnitts
Der Durchschnitt (`average`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-edcdfe1a4ae94847, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+300.00, -250.00, -23.00, +210.00, +234.00]
balance = analyse_transfers(transfers)
assert_equal(balance[1], 94.2, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+200.00, -200.00, +200.00, -200.00]
balance = analyse_transfers(transfers)
assert_equal(balance[1], 0, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [+15.00]
balance = analyse_transfers(transfers)
assert_equal(balance[1], 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 3 - Korrekte Bestimmung der höchsten Bankbewegung
Die höchste Bankbewegung (`highest_transfer`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-bbb9fd26de84aa07, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+300.00, -250.00, -23.00, +210.00, +234.00]
balance = analyse_transfers(transfers)
assert_equal(balance[2], 300, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+200.00, -200.00, +200.00, -200.00]
balance = analyse_transfers(transfers)
assert_equal(balance[2], 200, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [+15.00]
balance = analyse_transfers(transfers)
assert_equal(balance[2], 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

transfers = [+300.00, -500.00, -23.00, +210.00, +234.00]
balance = analyse_transfers(transfers)
assert_equal(balance[2], -500, "Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```