# Aufgabe 2

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Analyse von Materialeingängen und Materialabgängen implementiert werden. So soll aus einer großen Liste von Eingängen und Abgängen erkannt werden wie der neue Materialbestand ist. Außerdem soll die durchschnittliche Höhe der Materialbewegungen berechnet werden, sowie der größte Ein- oder Abgang bestimmt werden.

Implementieren Sie nun die Funktion `analyse_stock`. Diese Funktion erhält als Parameter eine Liste `incoming_outgoing`, welche in Form von Integer-Werten die Eingänge und Abgänge auf das Konto angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis soll der Materialbestand(`stock`) als Integer-Wert zurückgegeben werden.  
Die durchschnittliche Höhe(`average`) der Eingänge und Abgänge eines Materials.  
Die höchste Transaktion (`highest_transaction`). Hierbei sollen sowohl Abgänge als auch Eingänge berücksichtigt werden.  



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-3785c4bc8f420267, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def analyse_stock(incoming_outgoing):
    stock = 0
    highest_transaction = 0
    for transaction in incoming_outgoing:
        stock += transaction
        if transaction ** 2 > highest_transaction ** 2:
            highest_transaction = transaction
    
    average = stock / len(incoming_outgoing)
    return stock, average, highest_transaction
    
###END SOLUTION
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-5f14419ec5f6f74a, points: 0, grade: true, solution: false -->
```python
try:
    analyse_stock([+519, -250, -19])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung des Kontostands
Der Materialbestand (`stock`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-4947e67a6bdf03c9, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = analyse_stock(incoming_outgoing)
assert_equal(stock[0], 472, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = analyse_stock(incoming_outgoing)
assert_equal(stock[0], 88888889, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [-12, +20]
stock = analyse_stock(incoming_outgoing)
assert_equal(stock[0], 8, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 2 - Korrekte Berechnung des Durchschnitts
Der Durchschnitt (`average`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-a2c7901a696087f5, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+300, -15, +230, -123, +987]
average = analyse_stock(incoming_outgoing)
assert_equal(average[1], 275.8, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+200, -200, +200, -200]
average = analyse_stock(incoming_outgoing)
assert_equal(average[1], 0, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [+15]
average = analyse_stock(incoming_outgoing)
assert_equal(average[1], 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```

## Test 3 - Korrekte Bestimmung der höchsten Materialtransaktion
Die höchste Warenbewegung (`highest_transaction`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch Materialabgänge zu berücksichtigen.

<!-- grade_id: cell-32db5ee92c2544f0, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+100, -26, -23, +210, +234]
highest_transaction = analyse_stock(incoming_outgoing)
assert_equal(highest_transaction[2], 234, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+200, -200, +200, -200]
highest_transaction = analyse_stock(incoming_outgoing)
assert_equal(highest_transaction[2], 200, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [+15]
highest_transaction = analyse_stock(incoming_outgoing)
assert_equal(highest_transaction[2], 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

incoming_outgoing = [+300, -500, -23, +210, +234]
highest_transaction = analyse_stock(incoming_outgoing)
assert_equal(highest_transaction[2], -500, "Das Ergebnis ist falsch")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```