# Aufgabe 1

Es soll ein Programm zur automatischen Analyse von Kontobewegungen implementiert werden. So soll aus einer großen Liste von Zugängen und Abgängen erkannt werden welche die größte ist

Implementieren Sie nun die Funktion `analyse_transfers`. Diese Funktion erhält als Parameter eine Liste `transfers`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).   
Als Ergebnis soll der `höchste Transfer` als Float-Wert zurückgegeben werden ("highest_transfer"). Hierbei sollen sowohl negative als auch positive Transfers berücksichtigt werden.



### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-5186165662b3ed1f, points: 0, grade: false, solution: true -->
```python


    highest_transfer = 0
    for transfer in transfers:
        # prüfe ob der transfer höher ist als der aktuell höchste transfer und ändere gegebenenfalls den Wert von "highest_transfer"
    
  
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-e076cdd05aca88c1, points: 0, grade: true, solution: false -->
```python
try:
    analyse_transfers([+519.55, -250.00, -19.37])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bestimmung der höchsten Bankbewegung
Die höchste Bankbewegung (`highest_transfer`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-923ff467b69ba94d, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+100.45, -26.89, -23.00, +210.00, +234.00]
highest_transfer = analyse_transfers(transfers)
assert_equal(highest_transfer, 234, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

transfers = [+200.00, -200.00, +200.00, -200.00]
highest_transfer = analyse_transfers(transfers)
assert_equal(highest_transfer, 200, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

transfers = [+15.00]
highest_transfer = analyse_transfers(transfers)
assert_equal(highest_transfer, 15, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")

transfers = [+300.00, -500.00, -23.00, +210.00, +234.00]
highest_transfer = analyse_transfers(transfers)
assert_equal(highest_transfer, -500, "Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
print("Test 4 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
Test 4 erfolgreich
```