# Aufgabe 5

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Berechnung eines Materialbestandes, anhand von Eingängen und Abgängen eines Materials in einer Lagerhalle, implementiert werden.  
**Allerdings sollen Sie hierbei auf Schleifen verzichten. Lösen Sie diese Aufgabe durch eine rekursive Funktion.  
Das heißt die Funktion soll sich immer wieder selbst aufrufen mit einem veränderten Parameter.**

Implementieren Sie nun die Funktion `calculate_stock`. Diese Funktion erhält als Parameter eine Liste `incoming_outgoing`, welche in Form von Integer-Werten die Materialeingänge und Abgänge in die Lagerhalle angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis soll der Materialbestand in Form eines Integer-Wertes zurückgegeben werden.

### Hilfe
Sie können Ihre Lösung mit den automatisierten [Tests](#Automatische-Tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `For-Schleifen` in Python finden Sie [hier](../../50_loops/for_loop.ipynb).

<!-- grade_id: cell-642de40a4e8b0627, points: 0, grade: false, solution: true -->
```python
### Dies ist lediglich ein Pseudocode. Sie müssen den Code komplett umschreiben, die Logik sollte jedoch erhalten bleiben

def calculate_stock(incoming_outgoing):
    if (die Liste hat nur einen Eintrag):
        return (Eintrag aus der Liste)
    else:
        return (erster Eintrag aus Liste) + calculate_stock(Liste ab dem zweiten Element)
```

  ***
  # Automatische Tests
  Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-b36ad3269f19faf2, points: 0, grade: true, solution: false -->
```python
try:
    calculate_stock([-19, -12, +234])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Materialbestand (`stock`) soll für jede Liste korrekt berechnet werden.  
**Diese Testfälle können nicht überprüfen ob Sie die Aufgabe rekursiv gelöst haben. Achten Sie selbst darauf dass Sie keine Schleifen in Ihrem Programm nutzen.**

<!-- grade_id: cell-4c4ef5d3063aa0e9, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 472, "Das Ergebnis ist falsch")
print("Test 1 erfolgreich")

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 88888889, "Das Ergebnis ist falsch")
print("Test 2 erfolgreich")

incoming_outgoing = [-12, +20]
stock = calculate_stock(incoming_outgoing)
assert_equal(stock, 8, "Das Ergebnis ist falsch")
print("Test 3 erfolgreich")
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```